Homeostatic Synaptic Plasticity of Miniature Excitatory Postsynaptic Currents in Mouse Cortical Cultures Requires Neuronal Rab3A

Author response:

The following is the authors’ response to the previous reviews

General Response to Reviewers:

We thank the Reviewers for their comments, which continue to substantially improve the quality and clarity of the manuscript, and therefore help us to strengthen its message while acknowledging alternative explanations.

All three reviewers raised the concern that we have not proven that Rab3A is acting on a presynaptic mechanism to increase mEPSC amplitude after TTX treatment of mouse cortical cultures. The reviewers’ main point is that we have not shown a lack of upregulation of postsynaptic receptors in mouse cortical cultures. We want to stress that we agree that postsynaptic receptors are upregulated after activity block in neuronal cultures. However, the reviewers are not acknowledging that we have previously presented strong evidence at the mammalian NMJ that there is no increase in AChR after activity blockade, and therefore the requirement for Rab3A in the homeostatic increase in quantal amplitude points to a presynaptic contribution. We agree that we should restrict our firmest conclusions to the data in the current study, but in the Discussion we are proposing interpretations. We have added the following new text:

“The impetus for our current study was two previous studies in which we examined homeostatic regulation of quantal amplitude at the NMJ. An advantage of studying the NMJ is that synaptic ACh receptors are easily identified with fluorescently labeled alpha-bungarotoxin, which allows for very accurate quantification of postsynaptic receptor density. We were able to detect a known change due to mixing 2 colors of alpha-BTX to within 1% (Wang et al., 2005). Using this model synapse, we showed that there was no increase in synaptic AChRs after TTX treatment, whereas miniature endplate current increased 35% (Wang et al., 2005). We further showed that the presynaptic protein Rab3A was necessary for full upregulation of mEPC amplitude (Wang et al., 2011). These data strongly suggested Rab3A contributed to homeostatic upregulation of quantal amplitude via a presynaptic mechanism. With the current study showing that Rab3A is required for the homeostatic increase in mEPSC amplitude in cortical cultures, one interpretation is that in both situations, Rab3A is required for an increase in the presynaptic quantum.”

The point we are making is that the current manuscript is an extension of that work and interpretation of our findings regarding the variability of upregulation of postsynaptic receptors in our mouse cortical cultures further supports the idea that there is a Rab3Adependent presynaptic contribution to homeostatic increases in quantal amplitude.

Public Reviews:

Reviewer #1 (Public review):

Koesters and colleagues investigated the role of the small GTPase Rab3A in homeostatic scaling of miniature synaptic transmission in primary mouse cortical cultures using electrophysiology and immunohistochemistry. The major finding is that TTX incubation for 48 hours does not induce an increase in the amplitude of excitatory synaptic miniature events in neuronal cortical cultures derived from Rab3A KO and Rab3A Earlybird mutant mice. NASPM application had comparable effects on mEPSC amplitude in control and after TTX, implying that Ca2+-permeable glutamate receptors are unlikely modulated during synaptic scaling. Immunohistochemical analysis revealed no significant changes in GluA2 puncta size, intensity, and integral after TTX treatment in control and Rab3A KO cultures. Finally, they provide evidence that loss of Rab3A in neurons, but not astrocytes, blocks homeostatic scaling. Based on these data, the authors propose a model in which neuronal Rab3A is required for homeostatic scaling of synaptic transmission, potentially through GluA2-independent mechanisms.

The major finding - impaired homeostatic up-scaling after TTX treatment in Rab3A KO and Rab3 earlybird mutant neurons - is supported by data of high quality. However, the paper falls short of providing any evidence or direction regarding potential mechanisms. The data on GluA2 modulation after TTX incubation are likely statistically underpowered, and do not allow drawing solid conclusions, such as GluA2-independent mechanisms of up-scaling.

The study should be of interest to the field because it implicates a presynaptic molecule in homeostatic scaling, which is generally thought to involve postsynaptic neurotransmitter receptor modulation. However, it remains unclear how Rab3A participates in homeostatic plasticity.

Major (remaining) point:

(1) Direct quantitative comparison between electrophysiology and GluA2 imaging data is complicated by many factors, such as different signal-to-noise ratios. Hence, comparing the variability of the increase in mini amplitude vs. GluA2 fluorescence area is not valid. Thus, I recommend removing the sentence "We found that the increase in postsynaptic AMPAR levels was more variable than that of mEPSC amplitudes, suggesting other factors may contribute to the homeostatic increase in synaptic strength." from the abstract.

We have not removed the statement, but altered it to soften the conclusion. It now reads, “We found that the increase in postsynaptic AMPAR levels in wild type cultures was more variable than that of mEPSC amplitudes, which might be explained by a presynaptic contribution, but we cannot rule out variability in the measurement.”.

Similarly, the data do not directly support the conclusion of GluA2-independent mechanisms of homeostatic scaling. Statements like "We conclude that these data support the idea that there is another contributor to the TTX- induced increase in quantal size." should be thus revised or removed.

This particular statement is in the previous response to reviewers only, we deleted the sentence that starts, “The simplest explanation Rab3A regulates a presynaptic contributor….”. and “Imaging of immunofluorescence more variable…”. We deleted “ our data suggest….consistently leads to an increase in mEPSC amplitude and sometimes leads to….” We added “…the lack of a robust increase in receptor levels leaves open the possibility that there is a presynaptic contributor to quantal size in mouse cortical cultures. However, the variability could arise from technical factors associated with the immunofluorescence method, and the mechanism of Rab3A-dependent plasticity could be presynaptic for the NMJ and postsynaptic for cortical neurons.”

Reviewer #2 (Public review):

I thank the authors for their efforts in the revision. In general, I believe the main conclusion that Rab3A is required for TTX-induced homeostatic synaptic plasticity is wellsupported by the data presented, and this is an important addition to the repertoire of molecular players involved in homeostatic compensations. I also acknowledge that the authors are more cautious in making conclusions based on the current evidence, and the structure and logic have been much improved.

The only major concern I have still falls on the interpretation of the mismatch between GluA2 cluster size and mEPSC amplitude. The authors argue that they are only trying to say that changes in the cluster size are more variable than those in the mEPSC amplitude, and they provide multiple explanations for this mismatch. It seems incongruous to state that the simplest explanation is a presynaptic factor when you have all these alternative factors that very likely have contributed to the results. Further, the authors speculate in the discussion that Rab3A does not regulate postsynaptic GluA2 but instead regulates a presynaptic contributor. Do the authors mean that, in their model, the mEPSC amplitude increases can be attributed to two factors- postsynaptic GluA2 regulation and a presynaptic contribution (which is regulated by Rab3A)? If so, and Rab3A does not affect GluA2 whatsoever, shouldn't we see GluA2 increase even in the absence of Rab3A? The data in Table 1 seems to indicate otherwise.

The main body of this comment is addressed in the General Response to Reviewers. In addition, we deleted text “current data, coupled with our previous findings at the mouse neuromuscular junction, support the idea that there are additional sources contributing to the homeostatic increase in quantal size.” We added new text, so the sentence now reads: “Increased receptors likely contribute to increases in mESPC amplitudes in mouse cortical cultures, but because we do not have a significant increase in GluA2 receptors in our experiments, it is impossible to conclude that the increase is lacking in cultures from Rab3A-/- neurons.”

I also question the way the data are presented in Figure 5. The authors first compare 3 cultures and then 5 cultures altogether, if these experiments are all aimed to answer the same research question, then they should be pooled together. Interestingly, the additional two cultures both show increases in GluA2 clusters, which makes the decrease in culture #3 even more perplexing, for which the authors comment in line 261 that this is due to other factors. Shouldn't this be an indicator that something unusual has happened in this culture?

Data in this figure is sufficient to support that GluA2 increases are variable across cultures, which hardly adds anything new to the paper or to the field.

A major goal of performing the immunofluorescence measurements in the same cultures for which we had electrophysiological results was to address the common impression that the homeostatic effect itself is highly variable, as the reviewer notes in the comment “…GluA2 increases are variable across cultures…” Presumably, if GluA2 increases are the mechanism of the mEPSC amplitude increases, then variable GluA2 increases should correlate with variable mEPSC amplitude increases, but that is not what we observed. We are left with the explanation that the immunofluorescence method itself is very variable. We have added the point to the Discussion, which reads, “the variability could arise from technical factors associated with the immunofluorescence method, and the mechanism of Rab3A-dependent homeostatic plasticity could be presynaptic for the NMJ and postsynaptic for cortical neurons.”

Finally, the implication of “Shouldn’t this be an indicator that something unusual has happened in this culture?” if it is not due to culture to culture variability in the homeostatic response itself, is that there was a technical problem with accurately measuring receptor levels. We have no reason to suspect anything was amiss in this set of coverslips (the values for controls and for TTX-treated were not outside the range of values in other experiments). In any of the coverslips, there may be variability in the amount of primary anti-GluA2 antibody, as this was added directly to the culture rather than prepared as a diluted solution and added to all the coverslips. But to remove this one experiment because it did not give the expected result is to allow bias to direct our data selection.

The authors further cite a study with comparable sample sizes, which shows a similar mismatch based on p values (Xu and Pozzo-Miller 2007), yet the effect sizes in this study actually match quite well (both ~160%). P values cannot be used to show whether two effects match, but effect sizes can. Therefore, the statement in lines 411-413 "... consistently leads to an increase in mEPSC amplitudes, and sometimes leads to an increase in synaptic GluA2 receptor cluster size" is not very convincing, and can hardly be used to support "the idea that there are additional sources contributing to the homeostatic increase in quantal size.”

We have the same situation; our effect sizes match (19.7% increase for mEPSC amplitude; 18.1% increase for GluA2 receptor cluster size, see Table 1), but in our case, the p value for receptors does not reach statistical significance. Our point here is that there is published evidence that the variability in receptor measurements is greater than the variability in electrophysiological measurements. But we have softened this point, removing the sentences containing “…consistently leads and sometimes...” and “……additional sources contributing…”.

I would suggest simply showing mEPSC and immunostaining data from all cultures in this experiment as additional evidence for homeostatic synaptic plasticity in WT cultures, and leave out the argument for "mismatch". The presynaptic location of Rab3A is sufficient to speculate a presynaptic regulation of this form of homeostatic compensation.

We have removed all uses of the word “mismatch,” but feel the presentation of the 3 matched experiments, 23-24 cells (Figure 5A, D), and the additional 2 experiments for a total of 5 cultures, 48-49 cells (Figure 5C, F), is important in order to demonstrate that the lack of statistically significant receptor response is due neither to a variable homeostatic response in the mEPSC amplitudes, nor to a small number of cultures.

Minor concerns:

(1) Line 214, I see the authors cite literature to argue that GluA2 can form homomers and can conduct currents. While GluA2 subunits edited at the Q/R site (they are in nature) can form homomers with very low efficiency in exogenous systems such as HEK293 cells (as done in the cited studies), it's unlikely for this to happen in neurons (they can hardly traffic to synapses if possible at all).

We were unable to identify a key reference that characterized GluA2 homomers vs. heteromers in native cortical neurons, but we have rewritten the section in the manuscript to acknowledge the low conductance of homomers:

“…to assess whether GluA2 receptor expression, which will identify GluA2 homomers and GluA2 heteromers (the former unlikely to contribute to mEPSCs given their low conductance relative to heteromers (Swanson et al., 1997; Mansour et al., 2001)…”

(2) Lines 221-222, the authors may have misinterpreted the results in Turrigiano 1998. This study does not show that the increase in receptors is most dramatic in the apical dendrite, in fact, this is the only region they have tested. The results in Figures 3b-c show that the effect size is independent of the distance from soma.

Figure 3 in Turrigiano et al., shows that the increase in glutamate responsiveness is higher at the cell body than along the primary dendrite. We have revised our description to indicate that an increase in responsiveness on the primary dendrite has been demonstrated in Turrigiano et al. 1998.

“We focused on the primary dendrite of pyramidal neurons as a way to reduce variability that might arise from being at widely ranging distances from the cell body, or, from inadvertently sampling dendritic regions arising from inhibitory neurons. In addition, it has been shown that there is a clear increase in response to glutamate in this region (Turrigiano et al., 1998).”

“…synaptic receptors on the primary dendrite, where a clear increase in sensitivity to exogenously applied glutamate was demonstrated (see Figure 3 in (Turrigiano et al., 1998)).

(3) Lines 309-310 (and other places mentioning TNFa), the addition of TNFa to this experiment seems out of place. The authors have not performed any experiment to validate the presence/absence of TNFa in their system (citing only 1 study from another lab is insufficient). Although it's convincing that glia Rab3A is not required for homeostatic plasticity here, the data does not suggest Rab3A's role (or the lack of) for TNFa in this process.

We have modified the paragraph in the Discussion that addresses the glial results, to describe more clearly the data that supported an astrocytic TNF-alpha mechanism: “TNF-alpha accumulates after activity blockade, and directly applied to neuronal cultures, can cause an increase in GluA1 receptors, providing a potential mechanism by which activity blockade leads to the homeostatic upregulation of postsynaptic receptors (Beattie et al., 2002; Stellwagen et al., 2005; Stellwagen and Malenka, 2006).”

We have also acknowledged that we cannot rule out TNF-alpha coming from neurons in the cortical cultures: “…suggesting the possibility that neuronal Rab3A can act via a non-TNF-alpha mechanism to contribute to homeostatic regulation of quantal amplitude, although we have not ruled out a neuronal Rab3A-mediated TNF-alpha pathway in cortical cultures.”

Reviewer #3 (Public review):

This manuscript presents a number of interesting findings that have the potential to increase our understanding of the mechanism underlying homeostatic synaptic plasticity (HSP). The data broadly support that Rab3A plays a role in HSP, although the site and mechanism of action remain uncertain.

The authors clearly demonstrate that Rab3A plays a role in HSP at excitatory synapses, with substantially less plasticity occurring in the Rab3A KO neurons. There is also no apparent HSP in the Earlybird Rab3A mutation, although baseline synaptic strength is already elevated. In this context, it is unclear if the plasticity is absent, already induced by this mutation, or just occluded by a ceiling effect due to the synapses already being strengthened. Occlusion may also occur in the mixed cultures when Rab3A is missing from neurons but not astrocytes. The authors do appropriately discuss these options. The authors have solid data showing that Rab3A is unlikely to be active in astrocytes, Finally, they attempt to study the linkage between changes in synaptic strength and AMPA receptor trafficking during HSP, and conclude that trafficking may not be solely responsible for the changes in synaptic strength during HSP.

Strengths:

This work adds another player into the mechanisms underlying an important form of synaptic plasticity. The plasticity is likely only reduced, suggesting Rab3A is only partially required and perhaps multiple mechanisms contribute. The authors speculate about some possible novel mechanisms, including whether Rab3A is active pre-synaptically to regulate quantal amplitude.

As Rab3A is primarily known as a pre-synaptic molecule, this possibility is intriguing. However, it is based on the partial dissociation of AMPAR trafficking and synaptic response and lacks strong support. On average, they saw a similar magnitude of change in mEPSC amplitude and GluA2 cluster area and integral, but the GluA2 data was not significant due to higher variability. It is difficult to determine if this is due to biology or methodology - the imaging method involves assessing puncta pairs (GluA2/VGlut1) clearly associated with a MAP2 labeled dendrite. This is a small subset of synapses, with usually less than 20 synapses per neuron analyzed, which would be expected to be more variable than mEPSC recordings averaged across several hundred events. However, when they reduce the mEPSC number of events to similar numbers as the imaging, the mESPC amplitudes are still less variable than the imaging data. The reason for this remains unclear. The pool of sampled synapses is still different between the methods and recent data has shown that synapses have variable responses during HSP. Further, there could be variability in the subunit composition of newly inserted AMPARs, and only assessing GluA2 could mask this (see below). It is intriguing that pre-synaptic changes might contribute to HSP, especially given the likely localization of Rab3A. But it remains difficult to distinguish if the apparent difference in imaging and electrophysiology is a methodological issue rather than a biological one. Stronger data, especially positive data on changes in release, will be necessary to conclude that pre-synaptic factors are required for HSP, beyond the established changes in post-synaptic receptor trafficking.

Regarding the concern that the lack of increase in receptors is due to a technical issue, please see General Response to Reviewers, above. We have also softened our conclusions throughout, acknowledging we cannot rule out a technical issue.

Other questions arise from the NASPM experiments, used to justify looking at GluA2 (and not GluA1) in the immunostaining. First, there is a strong frequency effect that is unclear in origin. One would expect NASPM to merely block some fraction of the post-synaptic current, and not affect pre-synaptic release or block whole synapses. But the change in frequency seems to argue (as the authors do) that some synapses only have CP-AMPARs, while the rest of the synapses have few or none. Another possibility is that there are pre-synaptic NASPM-sensitive receptors that influence release probability. Further, the amplitude data show a strong trend towards smaller amplitude following NASPM treatment (Fig 3B). The p value for both control and TTX neurons was 0.08 - it is very difficult to argue that there is no effect. The decrease on average is larger in the TTX neurons, and some cells show a strong effect. It is possible there is some heterogeneity between neurons on whether GluA1/A2 heteromers or GluA1 homomers are added during HSP. This would impact the conclusions about the GluA2 imaging as compared to the mEPSC amplitude data.

The key finding in Figure 3 is that NASPM did not eliminate the statistically significant increase in mEPSC amplitude after TTX treatment (Fig 3A). Whether or not NASPM sensitive receptors contribute to mESPC amplitude is a separate question (Fig 3B). We are open to the possibility that NASPM reduces mEPSC amplitude in both control and TTX treated cells (p = 0.08 for both), but that does not change our conclusion that NASPM has no effect on the TTX-induced increase in mEPSC amplitude. The mechanism underlying the decrease in mEPSC frequency following NASPM is interesting, but does not alter our conclusions regarding the role of Rab3A in homeostatic synaptic plasticity of mEPSC amplitude. In addition, the Reviewer does not acknowledge the Supplemental Figure #1, which shows a similar lack of correspondence between homeostatic increases in mEPSC amplitude and GluA1 receptors in two cultures where matched data were obtained. Therefore, we do not think our lack of a robust increase in receptors can be explained by our failing to look at the relevant receptor.

To understand the role of Rab3A in HSP will require addressing two main issues:

(1) Is Rab3A acting pre-synaptically, post-synaptically or both? The authors provide good evidence that Rab3A is acting within neurons and not astrocytes. But where it is acting (pre or post) would aid substantially in understanding its role. The general view in the field has been that HSP is regulated post-synaptically via regulation of AMPAR trafficking, and considerable evidence supports this view. More concrete support for the authors' suggestion of a pre-synaptic site of control would be helpful.

We agree that definitive evidence for a presynaptic role of Rab3A in homeostatic plasticity of mEPSC amplitudes in mouse cortical cultures requires demonstrating that loss of Rab3A in postsynaptic neurons does not disrupt the plasticity, whereas loss in presynaptic neurons does. Without these data, we can only speculate that the Rab3A-dependence of homeostatic plasticity of quantal size in cortical neurons may be similar to that of the neuromuscular junction, where it cannot be receptors. We have added to the Discussion that the mechanism of Rab3A regulation of homeostatic plasticity of quantal amplitude could different between cortical neurons and the neuromuscular junction (lines 448-450 in markup,). Establishing a way to co-culture Rab3A-/- and Rab3A+/+ neurons in ratios that would allow us to record from a Rab3A-/- neuron that has mainly Rab3A+/+ inputs (or vice versa) is not impossible, but requires either transfection or transgenic expression with markers that identify the relevant genotype, and will be the subject of future experiments.

(2): Rab3A is also found at inhibitory synapses. It would be very informative to know if HSP at inhibitory synapses is similarly affected. This is particularly relevant as at inhibitory synapses, one expects a removal of GABARs or a decrease in GABA release (ie the opposite of whatever is happening at excitatory synapses). If both processes are regulated by Rab3A, this might suggest a role for this protein more upstream in the signaling; an effect only at excitatory synapses would argue for a more specific role just at those synapses.

We agree with the Reviewer, that it is important to determine the generality of Rab3A function in homeostatic plasticity. Establishing the homeostatic effect on mIPSCs and then examining them in Rab3A-/- cultures is a large undertaking and will be the subject of future experiments.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Minor (remaining) points:

(1) The figure referenced in the first response to the reviewers (Figure 5G) does not exist.

We meant Figure 5F, which has been corrected in the current response.

(2) I recommend showing the data without binning (despite some overlap).

The box plot in Origin will not allow not binning, but we can make the bin size so small that for all intents and purposes, there is close to 1 sample in each bin. When we do this, the majority of data are overlapped in a straight vertical line. Previously described concerns were regarding the gaps in the data, but it should be noted that these are cell means and we are not depicting the distributions of mEPSC amplitudes within a recording or across multiple recordings.

(3) Please auto-scale all axes from 0 (e.g., Fig 1E, F).

We have rescaled all mEPSC amplitude axes in box plots to go from 0 (Figures 1, 2 and 6).

(4) Typo in Figure legend 3: "NASPM (20 um)" => uM

Fixed.

Reviewer #2 (Recommendations for the authors):

(1) Line 140, frequencies are reported in Hz while other places are in sec-1, while these are essentially the same, they should be kept consistent in writing.

All mEPSC frequencies have been changed to sec-1, except we have left “Hz” for repetitive stimulation and filtering.

(2) Paragraph starting from line 163 (as well as other places where multiple groups are compared, such as the occlusion discussion), the authors assessed whether there was a change in baseline between WT and mutant group by doing pairwise tests, this is not the right test. A two-way ANOVA, or at least a multivariant test would be more appropriate.

We have performed a two-way ANOVA, with genotype as one factor, and treatment as the other factor. The p values in Figures 1 and 2 have been revised to reflect p values from the post-hoc Tukey test on the specific interactions (for each particular genotype, TTX vs CON effects). The difference in the two WT strains, untreated, was not significant in the Post-Hoc Tukey test, and we have revised the text. The difference between the untreated WT from the Rab3A+/Ebd colony and the untreated Rab3AEbd/Ebd mutant was still significant in the Post-Hoc Tukey test, and this has replaced the Kruskal-Wallis test. The two-way ANOVA was also applied to the neuron-glia experiments and p values in Figure 6 adjusted accordingly.

(3) Relevant to the second point under minor concerns, I suggest this sentence be removed, as reducing variability and avoiding inhibitory projects are reasons good enough to restrict the analysis to the apical dendrites.

We have revised the description of the Turrigiano et al., 1998 finding from their Figure 3 and feel it still strengthens the justification for choosing to analyze only synapses on the apical dendrite.

Reviewer #3 (Recommendations for the authors):

Minor points:

The comments on lines 256-7 could seem misleading - the NASPM results wouldn't rule out contribution of those other subunits, only non-GluA2 containing combinations of those subunits. I would suggest revising this statement. Also, NASPM does likely have an effect, just not one that changes much with TTX treatment.

At new line 213 (markup) we have added the modifier “homomeric” to clarify our point that the lack of NASPM effect on the increase in mEPSC amplitude after TTX indicates that the increase is not due to more homomeric Ca2+-permeable receptors. We have always stated that NASPM reduces mEPSC amplitude, but it is in both control and treated cultures.

Strong conclusions based on a single culture (lines 314-5) seem unwarranted.

We have softened this statement with a “suggesting that” substituted for the previous “Therefore,” but stand by our point that the mEPSC amplitude data support a homeostatic effect of TTX in Culture #3, so the lack of increase in GluA2 cluster size needs an explanation other than variability in the homeostatic effect itself.

Saying (line 554) something is 'the only remaining possibility' also seems unwarranted.

We have softened this statement to read, “A remaining possibility…”.

Beattie EC, Stellwagen D, Morishita W, Bresnahan JC, Ha BK, Von Zastrow M, Beattie MS, Malenka RC (2002) Control of synaptic strength by glial TNFalpha. Science 295:2282-2285.

Mansour M, Nagarajan N, Nehring RB, Clements JD, Rosenmund C (2001) Heteromeric AMPA receptors assemble with a preferred subunit stoichiometry and spatial arrangement. Neuron 32:841-853. Stellwagen D, Malenka RC (2006) Synaptic scaling mediated by glial TNF-alpha. Nature 440:1054-1059.

Stellwagen D, Beattie EC, Seo JY, Malenka RC (2005) Differential regulation of AMPA receptor and GABA receptor trafficking by tumor necrosis factor-alpha. J Neurosci 25:3219-3228.

Swanson GT, Kamboj SK, Cull-Candy SG (1997) Single-channel properties of recombinant AMPA receptors depend on RNA editing, splice variation, and subunit composition. J Neurosci 17:5869.

Turrigiano GG, Leslie KR, Desai NS, Rutherford LC, Nelson SB (1998) Activity-dependent scaling of quantal amplitude in neocortical neurons. Nature 391:892-896.

Wang X, Wang Q, Yang S, Bucan M, Rich MM, Engisch KL (2011) Impaired activity-dependent plasticity of quantal amplitude at the neuromuscular junction of Rab3A deletion and Rab3A earlybird mutant mice. J Neurosci 31:3580-3588.

Wang X, Li Y, Engisch KL, Nakanishi ST, Dodson SE, Miller GW, Cope TC, Pinter MJ, Rich MM (2005) Activity-dependent presynaptic regulation of quantal size at the mammalian neuromuscular junction in vivo. J Neurosci 25:343-351.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Following prolonged activity blockade, amplitudes of miniature excitatory postsynaptic currents (mEPSCs) increase, a form of plasticity termed “homeostatic synaptic plasticity.” We previously showed that a presynaptic protein, the small GTPase Rab3A, is required for full expression of the increase in miniature endplate current amplitudes following prolonged blockade of action potential activity at the mouse neuromuscular junction in vivo (Wang et al., 2011), but it is unknown whether this form of Rab3A-dependent homeostatic plasticity shares any characteristics with central synapses. We show here that homeostatic synaptic plasticity of mEPSCs is impaired in mouse cortical neuron cultures prepared from Rab3A-/- and mutant mice expressing a single point mutation of Rab3A, Rab3A Earlybird mice. To determine if Rab3A is involved in the well-established homeostatic increase in postsynaptic AMPA-type receptors (AMPARs), we performed a series of experiments in which electrophysiological recordings of mEPSCs and confocal imaging of synaptic AMPAR immunofluorescence were assessed within the same cultures. We found that the increase in postsynaptic AMPAR levels was more variable than that of mEPSC amplitudes, suggesting other factors may contribute to the homeostatic increase in synaptic strength. Finally, we demonstrate that Rab3A is acting in neurons because only selective loss of Rab3A in neurons, not glia, disrupted the homeostatic increase in mEPSC amplitudes. This is the first demonstration that a protein thought to function presynaptically is required for homeostatic synaptic plasticity of quantal size.

Revised version: Following prolonged activity blockade, amplitudes of miniature excitatory postsynaptic currents (mEPSCs) increase, a form of plasticity termed “homeostatic synaptic plasticity.” We previously showed that a presynaptic protein, the small GTPase Rab3A, is required for full expression of the increase in miniature endplate current amplitudes following prolonged blockade of action potential activity at the mouse neuromuscular junction in vivo, where an increase in postsynaptic receptors does not contribute (Wang et al., 2005; Wang et al., 2011). It is unknown whether this form of Rab3A-dependent homeostatic plasticity at the neuromuscular junction shares any characteristics with central synapses. We show here that homeostatic synaptic plasticity of mEPSCs is impaired in mouse cortical neuron cultures prepared from Rab3A-/- and mutant mice expressing a single point mutation of Rab3A, Rab3A Earlybird mice. To determine if Rab3A is involved in the well-established homeostatic increase in postsynaptic AMPA-type receptors (AMPARs), we performed a series of experiments in which electrophysiological recordings of mEPSCs and confocal imaging of synaptic AMPAR immunofluorescence were assessed within the same cultures. We found that the increase in postsynaptic AMPAR levels in wild type cultures was more variable than that of mEPSC amplitudes, which might be explained by a presynaptic contribution, but we cannot rule out variability in the measurement. Finally, we demonstrate that Rab3A is acting in neurons because only selective loss of Rab3A in neurons, not glia, disrupted the homeostatic increase in mEPSC amplitudes. This is the first demonstration that a protein thought to function presynaptically is required for homeostatic synaptic plasticity of quantal size in central neurons.

Introduction

One of the most studied phenomena triggered by prolonged activity blockade is the increase in amplitudes of miniature excitatory postsynaptic currents (mEPSCs) in neurons. First demonstrated in cultures of dissociated cortical neurons (Turrigiano et al., 1998) and spinal cord neurons (O’Brien et al., 1998), the compensatory response to prolonged loss of activity has been dubbed “homeostatic synaptic plasticity” (Turrigiano and Nelson, 2004; Pozo and Goda, 2010). These first two studies provided evidence that glutamate receptors were also increased after prolonged inactivity, via an increased response to exogenously applied glutamate (Turrigiano et al., 1998) and increased immunofluorescent labeling of GluA1 receptors at synaptic sites (O’Brien et al., 1998). The accompanying increase in synaptic AMPA receptors has been confirmed with immunofluorescence in multiple studies of hippocampal and cortical cultures treated with TTX, including (Wierenga et al., 2005; Hou et al., 2008; Ibata et al., 2008; Jakawich et al., 2010; Tatavarty et al., 2013; Xu and Pozzo-Miller, 2017; Dubes et al., 2022). It is now well accepted that the homeostatic increase in mEPSC amplitudes in neurons is due to an increase in postsynaptic AMPA receptors.

In our previous work, we found that a TTX cuff applied for 48 hrs around the sciatic nerve in mice led to an increase in the amplitude of the miniature endplate currents (mEPCs) recorded in tibialis anterior muscles. Surprisingly, and unlike at central synapses, we could find no evidence that the acetylcholine receptors levels at the neuromuscular junction (NMJ) were increased in the TTX-cuff treated mice (Wang et al., 2005). This result led us to search for presynaptic molecules that might homeostatically regulate the presynaptic quantum. In previous studies in chromaffin cells, we identified the small GTPase Rab3A, a presynaptic vesicle protein, as a regulator of synaptic vesicle fusion pore opening (Wang et al., 2008), so we examined whether deletion of Rab3A (Rab3A-/-) might prevent homeostatic upregulation of mEPC amplitude at the NMJ. The results were clear: the homeostatic increase in mEPC amplitude induced by the TTX cuff was strongly reduced in the Rab3A-/- mouse, and was completely abolished in the Earlybird mutant (Rab3AEbd/Ebd), which has a single point mutation in Rab3A that causes a shift towards early awakening due to a shorter circadian period, that is more dramatic than that of the Rab3A deletion mouse (Kapfhamer et al., 2002; Wang et al., 2011). These results led us to conclude that the homeostatic increase in mEPC amplitude at the NMJ is a presynaptic phenomenon.

It has remained somewhat of a mystery what the role of Rab3A is in synaptic transmission. The Rab3A-/- mouse has minimal phenotypic abnormalities, with evoked synaptic transmission and mEPSCs essentially normal in hippocampal slices (Geppert et al., 1994). At the Rab3A-/- NMJ, reductions in evoked transmission were detected, but only under conditions of reduced extracellular calcium (Coleman et al., 2007). There are some modest changes in short term plasticity during repetitive stimulus trains, with increased depression observed in response to moderate frequencies in hippocampal slices (Geppert et al., 1994) and at the mammalian NMJ (Coleman and Bykhovskaia, 2009), but increased facilitation in response to high frequency stimulation at the mammalian NMJ (Coleman and Bykhovskaia, 2010). The most dramatic effect of loss of Rab3A is the abolishment of a presynaptic form of long-term potentiation (LTP) at the mossy fiber-CA3 synapse that is induced by 25 Hz stimulation in the presence of NMDA blockers (Weisskopf et al., 1994; Castillo et al., 1997). Our result that prolonged inactivity-induced plasticity of mEPC amplitude at the mammalian NMJ is disrupted with loss of function of Rab3A further cements a role for Rab3A in activity-dependent plasticity of synaptic strength. However, a crucial question that remains is whether the TTX-cuff induced effect on the amplitude of the spontaneous synaptic current at the NMJ shares underlying mechanisms with central synapses. We set out to answer this question for the more well-studied phenomenon of homeostatic plasticity of mEPSCs in dissociated mouse cortical neurons prepared from wild type, Rab3A-/- mice and Rab3aEbd/Ebd mice treated with TTX for 48 hrs. We show here that homeostatic synaptic plasticity of mEPSC amplitude was strongly reduced in cultures from Rab3A-/- mice, and completely abolished in cultures from Rab3AEbd/Ebd mice, supporting a shared mechanism with the neuromuscular junction.

Results

Earlier version: We previously reported that mixed cultures of cortical neurons and glia prepared from postnatal day 0-2 mouse pups responded to a block of action potential-mediated activity by a 48 hr TTX treatment with an increase in mEPSC amplitude (Hanes et al., 2020). Here, we asked the question whether cortical cultures prepared from mice lacking the small GTPase Rab3A, or, expressing a point mutation of Rab3A, Rab3A Earlybird, have an altered homeostatic plasticity response to a prolonged loss of network activity. The Rab3A Earlybird mutation was discovered in a mutagenesis screen in mice for genes involved in the rest-activity cycle (Kapfhamer et al., 2002). In that study, the authors concluded that the Earlybird mutation was likely dominant-negative, as the shift in circadian period was substantially more dramatic than that observed in the Rab3A-/- mice. As described in the Introduction, the effects of loss of Rab3A have been modest, so we included experiments using the Rab3A Earlybird mutant mice to see if there would be a more robust phenotype, and in doing so, strengthen our conclusions based on the Rab3A deletion mice. To obtain Rab3A-/- and Rab3AEbd/Ebdhomozygotes, we established two mouse colonies of heterozygous breeders with cultures prepared from pups derived from a final breeding pair of homozygotes. Although we backcrossed Rab3A+/- with Rab3A+/Ebd for 11 generations, clear differences in mEPSC amplitudes in untreated cultures (see below) and in calcium current amplitudes in adrenal chromaffin cells (unpublished obs.) remained. Therefore, throughout this study we keep the two strains separate and there are two Rab3A+/+ or ‘wild type’ phenotypes.

Revised version: We previously reported that mixed cultures of cortical neurons and glia prepared from postnatal day 0-2 mouse pups responded to a block of action potential-mediated activity by a 48 hr TTX treatment with an increase in mEPSC amplitude (Hanes et al., 2020). Here, we asked the question whether cortical cultures prepared from mice lacking the small GTPase Rab3A, or, expressing a point mutation of Rab3A, Rab3A Earlybird, have an altered homeostatic plasticity response to a prolonged loss of network activity. The Rab3A Earlybird mutation was discovered in a mutagenesis screen in mice for genes involved in the rest-activity cycle (Kapfhamer et al., 2002). In that study, the authors concluded that the Earlybird mutation was likely dominant-negative, as the shift in circadian period was substantially more dramatic than that observed in the Rab3A-/- mice. As described in the Introduction, the effects of loss of Rab3A have been modest, so we included experiments using the Rab3A Earlybird mutant mice to see if there would be a more robust phenotype, and in doing so, strengthen our conclusions based on the Rab3A deletion mice. To obtain Rab3A-/- and Rab3AEbd/Ebd homozygotes, we established two mouse colonies of heterozygous breeders with cultures prepared from pups derived from a final breeding pair of homozygotes. Although we backcrossed Rab3A+/- with Rab3A+/Ebd for 11 generations, clear differences in mEPSC amplitudes in untreated cortical cultures for the two wild type strains (see below) and in calcium current amplitudes in adrenal chromaffin cells for the two wild type strains (unpublished obs.) remained. Therefore, throughout this study there are two Rab3A+/+ or ‘wild type’ phenotypes.

Earlier version: We found that loss of Rab3A resulted in a greatly impaired homeostatic increase after activity blockade. Example current traces of spontaneously occurring mEPSCs recorded from pyramidal neurons in untreated (CON) 13-14 DIV cortical cultures and sister cultures treated with 500 nM TTX for 48 hr prepared from wild type and Rab3A-/- mice in the Rab3A+/- colony are shown in Figures 1A and B, respectively. Average mEPSC waveforms from the same recordings are shown in Figures 1C and D. The mean mEPSC amplitudes for 30 control and 23 TTX-treated neurons from Rab3A+/+ cultures are displayed in the box and whisker plot in Figure 1E; after activity blockade the average mEPSC amplitude increased from 13.9 ± 0.7 pA to 18.2 ± 0.9 pA. In contrast, in cultures prepared from Rab3A-/- mice, the average mEPSC amplitude was not significantly increased, for 25 untreated cells and 26 TTX-treated cells (Figure 1F, 13.6 ± 0.1 pA vs. 14.3 ± 0.6 pA). We found that TTX treatment also resulted in an increase in mEPSC frequency in cultures prepared from Rab3A+/+ mice, as shown in Figure 1G (CON, 2.26 ± 0.37 sec-1; TTX, 4.62 ± 0.74 sec-1). This TTX-induced frequency increase was strongly reduced in cultures from Rab3A-/- mice, but there were a few outliers with very high frequencies still present after TTX treatment (Figure 1H, CON, 2.74 ± 0.49 Hz; TTX, 3.23 ± 0.93 Hz).

Revised version: We found that loss of Rab3A greatly impaired the homeostatic increase after activity blockade. Example current traces of spontaneously occurring mEPSCs recorded from pyramidal neurons in untreated (CON) 13-14 DIV cortical cultures and sister cultures treated with 500 nM TTX for 48 hr prepared from wild type and Rab3A-/- mice in the Rab3A+/- colony are shown in Figures 1A and B, respectively. Average mEPSC waveforms from the same recordings are shown in Figures 1C and D. The mean mEPSC amplitudes for 30 control and 23 TTX-treated neurons from Rab3A+/+ cultures are displayed in the box and whisker plot in Figure 1E; after activity blockade the average mEPSC amplitude increased from 13.9 ± 0.7 pA to 18.2 ± 0.9 pA. In contrast, in cultures prepared from Rab3A-/- mice, the average mEPSC amplitude was not significantly increased, for 25 untreated cells and 26 TTX-treated cells (Figure 1F, 13.6 ± 0.1 pA vs. 14.3 ± 0.6 pA). We found that TTX treatment also resulted in an increase in mEPSC frequency in cultures prepared from Rab3A+/+ mice, as shown in Figure 1G (CON, 2.26 ± 0.37 sec-1; TTX, 4.62 ± 0.74 sec-1). This TTX-induced frequency increase was strongly reduced in cultures from Rab3A-/- mice, but there were a few outliers with very high frequencies still present after TTX treatment (Figure 1H, CON, 2.74 ± 0.49 sec-1; TTX, 3.23 ± 0.93 sec-1).

Loss of Rab3A disrupted the TTX-induced increase in amplitudes of mEPSCs recorded in cultured mouse cortical neurons; the increase in frequency was more variable but also appears to be reduced.

Earlier version: (A) Ten second example traces recorded at −60 mV in pyramidal cortical neurons from an untreated (“CON”) and TTX-treated (“TTX”) neuron in cultures prepared from Rab3A+/+ mice from the Rab3A+/- colony. (B) Ten second example traces recorded at −60 mV in pyramidal cortical neurons from an untreated (“CON”) and TTX-treated (“TTX”) neuron in cultures prepared from Rab3A-/- mice. (C & D) Average traces for the recordings shown in A and B, respectively. (E) Box plots for average mEPSC amplitudes from untreated cells and TTX-treated cells in cultures prepared from Rab3A+/+ mice (Rab3A+/- colony; CON, N = 30 cells, 13.9 ± 0.7 pA; TTX, N = 23 cells, 18.2 ± 0.9 pA; from 11 cultures). (F) Box plots for average mEPSC amplitudes from untreated cells and TTX-treated cells in cultures prepared from Rab3A-/- mice (CON, N = 25 cells, 13.6 ± 0.7 pA; TTX, N = 26 cells, 14.3 ± 0.6 pA); from 11 cultures. (G) Box plots for average mEPSC frequency for same Rab3A+/+ cells as in (E) (CON, 2.26 ± 0.37 sec-1; TTX, 4.62 ± 0.74 sec- 1). (H) Box plots for average mEPSC frequency for same Rab3A-/-cells as in (F) (CON, 2.74 ± 0.49 sec-1; TTX, 3.23 ± 0.93 sec-1). Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10th and 90th percentiles; open circles represent means from individual cells, bin size 1 pA, 0.5 sec-1; line, median; dot, mean. p values (shown on the graphs) are from Kruskal-Wallis test. For all p-values, * with underline indicates significance with p < 0.05.

Revised version: (A) Ten second example traces recorded at -60 mV in pyramidal cortical neurons from an untreated (“CON”) and TTX-treated (“TTX”) neuron in cultures prepared from Rab3A+/+ mice from the Rab3A+/- colony. (B) Ten second example traces recorded at -60 mV in pyramidal cortical neurons from an untreated (“CON”) and TTX-treated (“TTX”) neuron in cultures prepared from Rab3A-/- mice. (C & D) Average traces for the recordings shown in A and B, respectively. (E) Box plots for average mEPSC amplitudes from untreated cells and TTX-treated cells in cultures prepared from Rab3A+/+ mice (Rab3A+/- colony; CON, N = 30 cells, 13.9 ± 0.7 pA; TTX, N = 23 cells, 18.2 ± 0.9 pA; from 11 cultures). (F) Box plots for average mEPSC amplitudes from untreated cells and TTX-treated cells in cultures prepared from Rab3A-/- mice (CON, N = 25 cells, 13.6 ± 0.7 pA; TTX, N = 26 cells, 14.3 ± 0.6 pA); from 11 cultures. (G) Box plots for average mEPSC frequency for same Rab3A+/+ cells as in (E) (CON, 2.26 ± 0.37 sec-1; TTX, 4.62 ± 0.74 sec- 1). (H) Box plots for average mEPSC frequency for same Rab3A-/-cells as in (F) (CON, 2.74 ± 0.49 sec-1; TTX, 3.23 ± 0.93 sec-1). Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10th and 90th percentiles; open circles represent means from individual cells, bin size 0.1 pA, 0.5 sec-1; line, median; dot, mean. p values (shown on the graphs) are from Tukey’s post-hoc test following a two-way ANOVA. For all p-values, * with underline indicates significance with p < 0.05.

Earlier version: We next examined homeostatic plasticity in cultures prepared from wild type mice in the Rab3A+/Ebd colony. Example current traces of mEPSCs, and average mEPSC waveforms, are shown in Figures 2A and C, respectively, for cortical cultures prepared from wild type mice in the Rab3A+/Ebd colony. Treatment with TTX for 48 hr led to a significant increase in the average mEPSC amplitude of 23 TTX-treated cells when compared to 20 untreated cells (Figure 2E, CON, 11.0 ± 0.6 pA; TTX, 15.0 ± 1.3 pA). We note here that while the two wild type strains respond very similarly to activity blockade, the mean mEPSC amplitude in untreated cultures was significantly different in the two wild type strains, 13.9 ± 0.7 pA, wild type from Rab3A+/- colony; 11.0 ± 0.6 pA, wild type from Rab3A+/Ebd colony (p = 0.004, Kruskal-Wallis test).

Revised version: We next examined homeostatic plasticity in cultures prepared from wild type mice in the Rab3A+/Ebd colony. Example current traces of mEPSCs, and average mEPSC waveforms, are shown in Figures 2A and C, respectively, for cortical cultures prepared from wild type mice in the Rab3A+/Ebd colony. Treatment with TTX for 48 hr led to a significant increase in the average mEPSC amplitude of 23 TTX-treated cells when compared to 20 untreated cells (Figure 2E, CON, 11.0 ± 0.6 pA; TTX, 15.1 ± 1.2 pA). We note here that while the two wild type strains respond very similarly to activity blockade, the mean mEPSC amplitude in untreated cultures was different between the two strains, 13.9 ± 0.7 pA, wild type from Rab3A+/- colony; 11.0 ± 0.6 pA, wild type from Rab3A+/Ebd colony, but this difference did not reach statistical significance in a Tukey means comparison after a two-way ANOVA (p = 0.27).

Normally functioning Rab3A was required for TTX-induced homeostatic plasticity of mEPSC amplitudes in cultured mouse cortical neurons.

Earlier version: (A) Ten second example traces recorded at −60 mV in pyramidal cortical neurons from an untreated (“CON”) and TTX-treated (“TTX”) neuron in cultures prepared from Rab3A+/+ mice from the Rab3A+/Ebd colony. (B) Ten second example traces recorded at −60 mV in pyramidal cortical neurons from an untreated (“CON”) and TTX-treated (“TTX”) neuron in cultures prepared from Rab3AEbd/Ebd mice. (C & D) Average traces for the recordings shown in A and B, respectively. (E) Box plots for average mEPSC amplitudes from untreated cells and TTX-treated cells from cultures prepared from Rab3A+/+ mice (Rab3A+/Ebd colony; CON, N = 20 cells, 11.0 ± 0.6 pA; TTX, N = 23 cells, 15.1 ± 1.2 pA; from 6 cultures). (F) Box plots for average mEPSC amplitudes from untreated cells and TTX-treated cells from cultures prepared from Rab3AEbd/Ebd mice (CON, N = 21 cells, 15.1 ± 1.0 pA; TTX, N = 22 cells, 14.6 ± 1.1 pA; from 7 cultures.) (G) Box plots for average mEPSC frequency for same Rab3A+/+ cells as in (E) (CON, 1.15 ± 0.19 sec-1; TTX, 2.54 ± 0.55 sec-1). (H) Box plots for average mEPSC frequency for same Rab3AEbd/Ebd cells as in (F) (CON, 1.71 ± 0.41 sec-1; TTX, 3.05 ± 0.80 sec-1). Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10th and 90th percentiles; open circles represent means from individual cells; bin size 1 pA, 0.5 sec-1; line, median; dot, mean. p values (shown on the graphs) are from Kruskal-Wallis test. For all p-values, * with underline indicates significance with p < 0.05.

Revised version: (A) Ten second example traces recorded at -60 mV in pyramidal cortical neurons from an untreated (“CON”) and TTX-treated (“TTX”) neuron in cultures prepared from Rab3A+/+ mice from the Rab3A+/Ebd colony. (B) Ten second example traces recorded at -60 mV in pyramidal cortical neurons from an untreated (“CON”) and TTX-treated (“TTX”) neuron in cultures prepared from Rab3AEbd/Ebd mice. (C & D) Average traces for the recordings shown in A and B, respectively. (E) Box plots for average mEPSC amplitudes from untreated cells and TTX-treated cells from cultures prepared from Rab3A+/+ mice (Rab3A+/Ebd colony; CON, N = 20 cells, 11.0 ± 0.6 pA; TTX, N = 23 cells, 15.1 ± 1.2 pA; from 6 cultures). (F) Box plots for average mEPSC amplitudes from untreated cells and TTX-treated cells from cultures prepared from Rab3AEbd/Ebd mice (CON, N = 21 cells, 15.1 ± 1.0 pA; TTX, N = 22 cells, 14.6 ± 1.1 pA; from 7 cultures.) (G) Box plots for average mEPSC frequency for same Rab3A+/+ cells as in (E) (CON, 1.15 ± 0.19 sec-1; TTX, 2.54 ± 0.55 sec-1). (H) Box plots for average mEPSC frequency for same Rab3AEbd/Ebd cells as in (F) (CON, 1.71 ± 0.41 sec-1; TTX, 3.05 ± 0.80 sec-1). Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10th and 90th percentiles; open circles represent means from individual cells; bin size 0.1 pA, 0.5 sec-1; line, median; dot, mean. p values (shown on the graphs) are from a Tukey’s post-hoc test following a two-way ANOVA. For all p-values, * with underline indicates significance with p < 0.05.

We found a complete disruption of homeostatic plasticity in cortical cultures prepared from Rab3AEbd/Ebd mice, as can been seen in viewing example mEPSC traces and average mEPSC waveforms (Figures 2B and D, respectively). The lack of TTX effect was confirmed in a comparison of mEPSC amplitude means for 21 untreated and 22 TTX-treated cells (Figure 2F, CON, 15.1 ± 1.0 pA vs. TTX, 14.6 ± 1.1 pA). There was a trend towards increased mEPSC frequency in TTX-treated cultures from this wild type strain, but the difference was not as robust as for the wild type cultures from Rab3A+/- colony, and did not reach statistical significance (Figure 2G, CON, 1.15 ± 0.19 sec-1; TTX, 2.54 ± 0.55 sec-1). This trend was not observed in neurons from cultures prepared from Rab3AEbd/Ebd mice, likely due to an increase in frequency in untreated cells—frequencies remained high in TTX-treated cells (Figure 2H, CON, 1.71 ± 0.41 sec-1; TTX, 3.05 ± 0.80 sec-1). Our results show that the homeostatic increase in mEPSC amplitude after activity blockade is disrupted in both Rab3A-/- and the Rab3AEbd/Ebdcortical neurons, strongly supporting a crucial role of functioning Rab3A in this process.

Earlier version: The disruption in homeostatic plasticity differs for Rab3A-/- and Rab3AEbd/Ebd. In the Rab3A-/- data set, mEPSCs from untreated cultures were indistinguishable from mEPSCs from Rab3A+/+ untreated cultures, demonstrating loss of Rab3A has no impact on basal activity mEPSC amplitudes. In the Rab3AEbd/Ebd data set, mEPSC amplitudes from untreated cultures were significantly larger than those of untreated cultures from wild type mice, as can be seen comparing the CON values in Figures 2E and F (CON, Rab3A+/+, 11.0 ± 0.7 pA, vs. CON, Rab3AEbd/Ebd, 15.1 ± 1.0 pA, p = 0.0027, Kruskal-Wallis test). The increase in mEPSC amplitude in cultures from Rab3AEbd/Ebd mice is consistent with the increase in mEPC amplitude we observed at the Rab3AEbd/EbdNMJ (Wang et al., 2011).

Revised version: The disruption in homeostatic plasticity differs for Rab3A-/- and Rab3AEbd/Ebd. In the Rab3A-/- data set, mEPSCs from untreated cultures were indistinguishable from mEPSCs from Rab3A+/+ untreated cultures, demonstrating loss of Rab3A has no impact on basal activity mEPSC amplitudes. In the Rab3AEbd/Ebddata set, mEPSC amplitudes from untreated cultures were significantly larger than those of untreated cultures from wild type mice, as can be seen comparing the CON values in Figures 2E and F (CON, Rab3A+/+, 11.0 ± 0.7 pA, vs. CON, Rab3AEbd/Ebd, 15.1 ± 1.0 pA, p = 0.04, Tukey means comparison after two-way ANOVA). The increase in mEPSC amplitude in cultures from Rab3AEbd/Ebd mice is consistent with the increase in mEPC amplitude we observed at the Rab3AEbd/Ebd NMJ (Wang et al., 2011).

The small GTPase Rab3A is generally thought to function presynaptically to regulate synaptic vesicle trafficking, possibly in an activity-dependent manner (Castillo et al., 1997; Lonart et al., 1998; Leenders et al., 2001; Schluter et al., 2006; Coleman and Bykhovskaia, 2009; Tian et al., 2012). In contrast, homeostatic plasticity of mEPSC amplitude has been attributed to an increase in postsynaptic receptors on the surface of the dendrite (O’Brien et al., 1998; Turrigiano et al., 1998). At the NMJ in vivo, we could find no evidence for an increase in AChRs after TTX block of the sciatic nerve in vivo (Wang et al., 2005). Is Rab3A required for the homeostatic increase in surface AMPA-type glutamate receptors that has been confirmed by multiple studies (see Introduction)?

Earlier version: We explored whether Rab3A might be acting to increase postsynaptic AMPA receptors. In some studies of homeostatic synaptic plasticity it has been found that application of a Ca2+- permeable AMPA receptor-specific inhibitor reversed the homeostatic increase in mEPSC amplitude (Ju et al., 2004; Thiagarajan et al., 2005; Sutton et al., 2006). Figure 3A shows that the TTX-induced increase in mean mEPSC amplitude was nearly identical in a set of 11 CON and 11 TTX-treated cells before and after treatment with 1-Naphthyl acetyl spermine trihydrochloride (NASPM), a synthetic analogue of the Ca2+-permeable AMPA receptor inhibitor philanthotoxin (before NASPM, CON 12.9 ± 3.5 pA; TTX, 17.5 ± 3.1 pA; after NASPM, CON 11.9 ± 2.6 pA; TTX 16.1 ± 3.5 pA). This result indicates that the TTX-induced increase in mEPSC amplitude does not depend on Ca2+-permeable receptors, since the effect of TTX clearly remained when their presence was removed by acute NASPM application. We do not think that there was any technical issue with the NASPM application, because overall, mEPSC amplitudes were reduced modestly in both untreated and TTX-treated cultures (Figure 3B, CON, before NASPM, 12.9 ± 3.5 pA; after NASPM, 11.9 ± 2.6 pA; TTX, before NASPM, 17.5 ± 3.1 pA; after NASPM, 16.1 ± 3.5 pA). Furthermore, NASPM consistently reduced mEPSC frequency (Figure 3C, CON, before NASPM, 1.84 ± 0.55 sec-1; after NASPM, 1.56 ± 0.53 sec-1; TTX, before NASPM, 4.40 ± 3.51 sec-1; after NASPM, 2.68 ± 2.25 sec-1). Similar to the data in Figure 1, frequency is significantly increased after TTX treatment (p = 0.04), but this difference is no longer significant after NASPM application (p = 0.20). One explanation for the effect on frequency is that there are synaptic sites that express only homomers of Ca2+-permeable AMPA receptors (GluA2-lacking AMPARs); mEPSCs from these sites would be expected to be completely blocked by NASPM (see cartoon description in Figure 3D, Left). We cannot rule out an alternative explanation, that presynaptic Ca2+-permeable AMPA receptors normally enhance release probability, and NAPSM prevents this action. The magnitude of the frequency reduction following acute NASPM appears to be greater after TTX treatment, suggesting that loss of activity could promote the establishment of synaptic sites that contain only Ca2+-permeable receptors. However, these new sites do not appear to contribute to the increase in mEPSC amplitude (Figure 3A). The lack of effect of NASPM rules out the contribution of Ca2+-permeable AMPA receptors GluA1, GluA3 and GluA4 (Hollmann et al., 1991; Burnashev et al., 1992; Jonas and Burnashev, 1995), as well as Ca2+-permeable forms of Kainate receptors GluA5 and GluA6 (Koike et al., 1997; Sun et al., 2009) to homeostatic synaptic plasticity.

Revised version: In some studies of homeostatic synaptic plasticity it has been found that application of a Ca2+-permeable AMPA receptor-specific inhibitor reversed the homeostatic increase in mEPSC amplitude (Ju et al., 2004; Thiagarajan et al., 2005; Sutton et al., 2006). Figure 3A shows that the TTX-induced increase in mean mEPSC amplitude was nearly identical in a set of 11 CON and 11 TTX-treated cells before and after treatment with 1-Naphthyl acetyl spermine trihydrochloride (NASPM), a synthetic analogue of the Ca2+-permeable AMPA receptor inhibitor philanthotoxin (before NASPM, CON 12.9 ± 3.5 pA; TTX, 17.5 ± 3.1 pA; after NASPM, CON 11.9 ± 2.6 pA; TTX 16.1 ± 3.5 pA). This result indicates that the TTX-induced increase in mEPSC amplitude does not depend on Ca2+-permeable receptors, since the effect of TTX clearly remained when their presence was removed by acute NASPM application. We do not think that there was any technical issue with the NASPM application, because overall, mEPSC amplitudes were reduced modestly in both untreated and TTX-treated cultures (Figure 3B, CON, before NASPM, 12.9 ± 3.5 pA; after NASPM, 11.9 ± 2.6 pA; TTX, before NASPM, 17.5 ± 3.1 pA; after NASPM, 16.1 ± 3.5 pA). Furthermore, NASPM consistently reduced mEPSC frequency (Figure 3C, CON, before NASPM, 1.84 ± 0.55 sec-1; after NASPM, 1.56 ± 0.53 sec-1; TTX, before NASPM, 4.40 ± 3.51 sec-1; after NASPM, 2.68 ± 2.25 sec-1). Similar to the data in Figure 1, frequency is significantly increased after TTX treatment (p = 0.04), but this difference is no longer significant after NASPM application (p = 0.20). One explanation for the effect on frequency is that there are synaptic sites that express only homomers of Ca2+-permeable AMPA receptors (GluA2-lacking AMPARs); mEPSCs from these sites would be expected to be completely blocked by NASPM (see cartoon description in Figure 3D, Left). We cannot rule out an alternative explanation, that presynaptic Ca2+-permeable AMPA receptors normally enhance release probability, and NAPSM prevents this action. The magnitude of the frequency reduction following acute NASPM appears to be greater after TTX treatment, suggesting that loss of activity could promote the establishment of synaptic sites that contain only Ca2+-permeable receptors. However, these new sites do not appear to contribute to the increase in mEPSC amplitude (Figure 3A). The lack of effect of NASPM rules out the contribution of homomeric Ca2+-permeable AMPA receptors GluA1, GluA3 and GluA4 (Hollmann et al., 1991; Burnashev et al., 1992; Jonas and Burnashev, 1995), as well as Ca2+-permeable forms of Kainate receptors GluA5 and GluA6 (Koike et al., 1997; Sun et al., 2009) to homeostatic synaptic plasticity.

Homeostatic plasticity of mEPSC amplitudes in mouse cortical cultures treated with TTX for 48 hr was unchanged by acute inhibition of Ca2+ permeable AMPA receptors by NASPM (20 μm).

Earlier version: (A) Box plot comparison of the TTX effect on mEPSC amplitudes in the same pyramidal neurons before and after application of 20 μM NASPM. (N = 11 cells from 3 cultures; Pre-NASPM, CON: 12.9 ± 1.1 pA; TTX: 17.5 ± 0.9 pA; post-NASPM, CON: 11.9 ± 0.8 pA; TTX: 16.1 ± 1.0 pA). (B) Line series plot of mEPSC amplitudes before and after acute perfusion with 20 μM NASPM for untreated and TTX-treated pyramidal neurons; same cells as in (A); CON, pre-NASPM: 12.9 ± 1.1 pA; post-NASPM: 11.9 ± 0.8 pA; TTX, pre-NASPM: 17.5 ± 0.9 pA; post-NASPM: 16.1 ± 0.1 pA. (C) Line series plot of mEPSC frequency before and after acute perfusion with 20 μM NASPM for untreated and TTX-treated pyramidal neurons; same cells as in (A); CON, pre-NAPSM: 1.84 ± 0.55 sec-1; post-NASPM; 1.56 ± 0.53 sec-1; TTX, pre-NASPM: 4.40 ± 1.06 sec-1; post-NASPM, 2.68 ± 0.68 sec-1. (D) A proposed mechanism for why NASPM had a robust effect on frequency without greatly affecting amplitude. A dendrite with spines (Top) is expanded on three postsynaptic sites (Middle) to show possible types of AMPA receptor distributions: Left, at a site comprised only of Ca2+-permeable AMPA receptors (NASPM-sensitive, GluA2-lacking receptors (black)), NASPM would completely inhibit the mEPSC, causing a decrease in frequency to be measured in the overall population; Right, at a site comprised only of Ca2+-impermeable AMPA receptors (NASPM-insensitive, GluA2-containing receptors (green)), NASPM would have no effect on mEPSC amplitude; Middle, at a site comprised of a mix of Ca2+-permeable and Ca2+- impermeable AMPA receptors, NASPM would partially inhibit the mEPSC. Diagram produced in BioRender.com (2024). Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10th and 90th percentiles; open circles represent means from individual cells; bin size 1 pA; line, median; dot, mean. p values (shown on the graphs) are from Kruskal-Wallis test. Line series plot p values are from student’s paired t test. For all p-values, * with underline indicates significance with p < 0.05.

Revised version: (A) Box plot comparison of the TTX effect on mEPSC amplitudes in the same pyramidal neurons before and after application of 20 μM NASPM. (N = 11 cells from 3 cultures; Pre-NASPM, CON: 12.9 ± 1.1 pA; TTX: 17.5 ± 0.9 pA; post-NASPM, CON: 11.9 ± 0.8 pA; TTX: 16.1 ± 1.0 pA). (B) Line series plot of mEPSC amplitudes before and after acute perfusion with 20 μM NASPM for untreated and TTX-treated pyramidal neurons; same cells as in (A); CON, pre-NASPM: 12.9 ± 1.1 pA; post-NASPM: 11.9 ± 0.8 pA; TTX, pre-NASPM: 17.5 ± 0.9 pA; post-NASPM: 16.1 ± 0.1 pA. (C) Line series plot of mEPSC frequency before and after acute perfusion with 20 μM NASPM for untreated and TTX-treated pyramidal neurons; same cells as in (A); CON, pre-NAPSM: 1.84 ± 0.55 sec-1; post-NASPM; 1.56 ± 0.53 sec-1; TTX, pre-NASPM: 4.40 ± 1.06 sec-1; post-NASPM, 2.68 ± 0.68 sec-1. (D) A proposed mechanism for why NASPM had a robust effect on frequency without greatly affecting amplitude. A dendrite with spines (Top) is expanded on three postsynaptic sites (Middle) to show possible types of AMPA receptor distributions: Left, at a site comprised only of Ca2+-permeable AMPA receptors (NASPM-sensitive, GluA2-lacking receptors (black)), NASPM would completely inhibit the mEPSC, causing a decrease in frequency to be measured in the overall population; Right, at a site comprised only of Ca2+-impermeable AMPA receptors (NASPM-insensitive, GluA2-containing receptors (green)), NASPM would have no effect on mEPSC amplitude; Middle, at a site comprised of a mix of Ca2+-permeable and Ca2+-impermeable AMPA receptors, NASPM would partially inhibit the mEPSC. Diagram produced in BioRender.com (2024). Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10th and 90th percentiles; open circles represent means from individual cells; bin size 0.1 pA; line, median; dot, mean. p values (shown on the graphs) are from Kruskal-Wallis test. Line series plot p values are from student’s paired t test. For all p-values, * with underline indicates significance with p < 0.05.

Earlier version: Having established that homomeric Ca2+-permeable receptors are not contributing to the homeostatic increase in mEPSC amplitude, we turned to immunofluorescence and confocal imaging to assess whether GluA2 receptor expression, which will identify GluA2 homomers and GluA2 heteromers, was increased in our wild type mouse cortical cultures following 48 hr treatment with TTX. Since mEPSCs necessarily report synaptic levels of receptors, we used VGLUT1-immunofluorescence to identify synapses on pyramidal primary apical dendrites labeled with MAP-2 immunofluorescence. We focused on the primary dendrite of pyramidal neurons as a way to reduce variability that might arise from being at widely ranging distances from the cell body, or, from inadvertently sampling dendritic regions arising from inhibitory neurons. In addition, it has been shown that the increase in receptors is most dramatic in this region (Turrigiano et al., 1998). Figure 4 (Top) shows two pyramidal neurons, one from an untreated culture on the left, and one from a TTX-treated culture on the right, prepared from Rab3A+/+ mice. To measure characteristics of synaptically located GluA2 receptor clusters, we zoomed in on the primary dendrite (region in white rectangle). The zoomed regions for single confocal sections of the selected area are shown below and contain pairs of VGLUT1- and GluA2-immunofluorescent clusters indicated by white trapezoidal frames. A dendrite typically required ∼10 confocal sections to be fully captured, and the total number of synaptic pairs for all the sections imaged on a dendrite averaged 20 (20.4 ± 6.5 (SD); range, 11-38) so this is an atypically high number of pairs within a single section; these particular dendrites and sections were selected for illustration purposes. In addition to the synaptic pairs, we observed many GluA2-immunofluorescent clusters not associated with VGLUT1 immunofluorescence (a few are indicated with white arrows), most probably extrasynaptic receptors. There were also GluA2- positive clusters present outside of MAP2-positive dendrites, which may be located on astrocytes (Fan et al., 1999). Finally, GluA2 immunofluorescent clusters close to VGLUT1 immunofluorescence but not located along any apparent MAP2-positive neurites suggests the presence of axon-axonal contacts, although VGLUT1 has also been detected in astrocytes (Ormel et al., 2012). Only sites that contained both VGLUT1 and GluA2 immunofluorescence close to the primary MAP2-postive dendrite or a secondary branch along this primary dendrite were selected for analyses. In the experiment from which these images were collected, we analyzed the distance from the cell body for each synapse. The average distance from the cell body, for dendrites from the untreated cell, was 38.5 ± 2.8 μm; for the TTX-treated cell, it was 42.4 ± 3.2 μm (p = 0.35, Kruskal-Wallis test). Since the greatest distance from the cell body is set by the outer limit of the zoom window, and that window was placed adjacent to the cell body, distances in other experiments would be in this same range.

Revised version: Having established that homomeric Ca2+-permeable receptors are not contributing to the homeostatic increase in mEPSC amplitude, we turned to immunofluorescence and confocal imaging to assess whether GluA2 receptor expression, which will identify GluA2 homomers and GluA2 heteromers (the former unlikely to contribute to mEPSCs given their low conductance relative to heteromers (Swanson et al., 1997; Mansour et al., 2001), was increased in our wild type mouse cortical cultures following 48 hr treatment with TTX. Since mEPSCs necessarily report synaptic levels of receptors, we used VGLUT1-immunofluorescence to identify synapses on pyramidal primary apical dendrites labeled with MAP-2 immunofluorescence. We focused on the primary dendrite of pyramidal neurons as a way to reduce variability that might arise from being at widely ranging distances from the cell body, or, from inadvertently sampling dendritic regions arising from inhibitory neurons. In addition, it has been shown that there is a clear increase in response to glutamate in this region (Turrigiano et al., 1998). Figure 4 (Top) shows two pyramidal neurons, one from an untreated culture on the left, and one from a TTX-treated culture on the right, prepared from Rab3A+/+ mice. To measure characteristics of synaptically located GluA2 receptor clusters, we zoomed in on the primary dendrite (region in white rectangle). The zoomed regions for single confocal sections of the selected area are shown below and contain pairs of VGLUT1- and GluA2-immunofluorescent clusters indicated by white trapezoidal frames. A dendrite typically required ∼10 confocal sections to be fully captured, and the total number of synaptic pairs for all the sections imaged on a dendrite averaged 20 (20.4 ± 6.5 (SD); range, 11-38) so this is an atypically high number of pairs within a single section; these particular dendrites and sections were selected for illustration purposes. In addition to the synaptic pairs, we observed many GluA2-immunofluorescent clusters not associated with VGLUT1 immunofluorescence (a few are indicated with white arrows), most probably extrasynaptic receptors. There were also GluA2-positive clusters present outside of MAP2-positive dendrites, which may be located on astrocytes (Fan et al., 1999). Finally, GluA2 immunofluorescent clusters close to VGLUT1 immunofluorescence but not located along any apparent MAP2-positive neurites suggests the presence of axon-axonal contacts, although VGLUT1 has also been detected in astrocytes (Ormel et al., 2012). Only sites that contained both VGLUT1 and GluA2 immunofluorescence close to the primary MAP2-postive dendrite or a secondary branch along this primary dendrite were selected for analyses. In the experiment from which these images were collected, we analyzed the distance from the cell body for each synapse. The average distance from the cell body, for dendrites from the untreated cell, was 38.5 ± 2.8 μm; for the TTX-treated cell, it was 42.4 ± 3.2 μm (p = 0.35, Kruskal-Wallis test). Since the greatest distance from the cell body is set by the outer limit of the zoom window, and that window was placed adjacent to the cell body, distances in other experiments would be in this same range.

Identification of synaptic GluA2 receptor immunofluorescence on primary dendrites of pyramidal neurons in high density mouse cortical co-cultures prepared from Rab3A+/+ mice.

(Top) Non-zoomed single confocal sections collected with a 60X oil immersion objective of recognizably pyramidal-shaped neurons, presumed to be excitatory neurons, selected for synaptic GluA2 analysis. A neuron was selected from an untreated coverslip (CON, Left) and a TTX-treated coverslip (TTX, Right) from the same culture prepared from Rab3A+/+ animals (= Rab3A+/+ Culture #2 in Figure 5). The white rectangular boxes indicate 5X zoomed areas shown in the images below. Note that the depth of the single confocal section of the non-zoomed neuron image is not at the same depth as the confocal section in zoomed dendritic image, so some features are not visible in both images. Scale bar, 20 μm. (Bottom) 5X zoom single confocal sections selected for demonstration purposes because they had an unusually high number of identified synaptic pairs along the primary dendrite contained within a single confocal section. Synaptic pairs, highlighted with white trapezoids, were identified based on close proximity of GluA2 (red) and VGLUT1 (white) immunofluorescence, apposed to the MAP2 immunofluorescent primary dendrite (green). Some apparent synaptic pairs are not highlighted with a white trapezoid because there was a different confocal section in which the two immunofluorescent sites were maximally bright. There were GluA2 positive clusters located on the primary dendrites that were not apposed to VGLUT1-positive terminals (4 of these non-synaptic GluA2 clusters are highlighted with white arrows in the MAP2-GluA2 and MAP2-GluA2-VGLUT1 panels). Scale bar, 10 μm.

Earlier version: Variability in the magnitude of the homeostatic response from culture to culture is averaged out in physiological experiments by the pooling of data from many cells across many cultures. To reduce the necessity for many cultures, we chose to pair experiments in the same cultures by recording mEPSCs from one set of coverslips, and processing another set of coverslips from the same culture for immunofluorescence. We also wanted to determine if the homeostatic response of mEPSC amplitudes varied from culture to culture, and if so, whether the GluA2 receptor levels varied in parallel. We completed this matched paradigm of physiology and immunofluorescence for 3 cultures prepared from Rab3A+/+ mice. We first present the results as means for each culture, for control and TTX-treated dishes, for mEPSC data (Figure 5A) and imaging data (Figure 5B). Levels of GluA2 immunofluorescence at synaptic sites were quantified by the size of the GluA2-immunofluorescent receptor cluster, because this showed the greatest response to activity blockade; the average intensity value, and the integral of intensity values across the area of the cluster, are included in a summary of the data in Table 1. Three of three cortical cultures showed an increase in mean mEPSC amplitude after TTX treatment. In contrast, one of the same three cultures showed a decrease in mean GluA2 receptor cluster size. Therefore, the decrease in receptor levels in Culture #3 was due to something other than a lack of homeostatic effect on mEPSCs in that culture. Because there was a clear trend towards increased GluA2 receptor cluster size, but the receptor response appears more variable, we performed immunofluorescence measurements on two additional cortical cultures, both of which showed modest increases after TTX treatment (Figure 5C). We confirmed that there was not an increase in GluA1 receptor levels in two matched cultures that showed robust homeostatic increases in mEPSC amplitudes (Supplementary Figure 1A, culture mean mEPSC amplitudes; Supplementary Figure 1B, culture mean GluA1 receptor cluster sizes).

Revised version: Variability in the magnitude of the homeostatic response from culture to culture is averaged out in physiological experiments by the pooling of data from many cells across many cultures. To reduce the necessity for many cultures, we chose to pair experiments in the same cultures by recording mEPSCs from one set of coverslips, and processing another set of coverslips from the same culture for immunofluorescence. We also wanted to determine if the homeostatic response of mEPSC amplitudes varied from culture to culture, and if so, whether the GluA2 receptor levels varied in parallel. We completed this matched paradigm of physiology and immunofluorescence for 3 cultures prepared from Rab3A+/+ mice. We first present the results as means for each culture, for control and TTX-treated dishes, for mEPSC data (Figure 5A) and imaging data (Figure 5B). Levels of GluA2 immunofluorescence at synaptic sites were quantified by the size of the GluA2-immunofluorescent receptor cluster, because this showed the greatest response to activity blockade; the average intensity value, and the integral of intensity values across the area of the cluster, are included in a summary of the data in Table 1. Three of three cortical cultures showed an increase in mean mEPSC amplitude after TTX treatment. In contrast, one of the same three cultures showed a decrease in mean GluA2 receptor cluster size, suggesting that the decrease in receptor levels in Culture #3 was not due to a lack of homeostatic effect on mEPSC amplitudes in that culture. Because there was a clear trend towards increased GluA2 receptor cluster size, but the receptor response appears more variable, we performed immunofluorescence measurements on two additional cortical cultures, both of which showed modest increases after TTX treatment (Figure 5C). We confirmed that there was also not a consistent increase in GluA1 receptor levels in two matched cultures that showed robust homeostatic increases in mEPSC amplitudes (Supplementary Figure 1A, culture mean mEPSC amplitudes; Supplementary Figure 1B, culture mean GluA1 receptor cluster sizes).

Comparison of mEPSC amplitudes and GluA2 receptor cluster areas in matched and unmatched mouse cortical cultures prepared from Rab3A+/+ mice and treated with TTX for 48 hr.

Earlier version: (A) Culture averages of mEPSC amplitudes for untreated (CON) and TTX-treated coverslips (TTX) in each of 3 Rab3A+/+ mouse cortical co-cultures. Culture #1, CON, N = 6, 14.2 ± 2.2 pA; TTX, N = 6, 15.9 ± 1.9 pA; Culture #2, CON, N = 7, 13.8 ± 2.4 pA; TTX, N = 8, 17.7 ± 1.8 pA; Culture #3, CON, N = 10, 13.4 ± 0.8 pA; TTX, N = 9, 15.7 ± 1.0 pA. (B) Culture averages of GluA2 receptor cluster size in the same 3 cultures as shown in (A). Culture #1, CON, N = 10, 0.73 ± 0.09 μm2; TTX, N = 9, 0.89 ± 0.09 μm2; Culture #2, CON, N = 9, 0.91 ± 0.12 μm2; TTX, N = 9, 1.42 ± 0.24 μm2; Culture #3, CON, N = 10, 1.05 ± 0.11 μm2; TTX, N = 9, 0.95 ± 0.15 μm2. (C) Two additional culture averages are included that did not have corresponding mEPSC recordings. Culture #4, CON, N = 10, 0.58 ± 0.07 μm2: TTX, N = 10, 0.71 ± 0.04 μm2; Culture 5, CON, N =10, 0.89 ± 0.05 μm2; TTX, N = 10, 0.95 ± 0.08 μm2. (D) Box plots of cell mean mEPSC amplitudes pooled from same 3 cultures as shown in (A). CON, N = 23 cells, 13.7 ± 0.9 pA; TTX, N = 24 cells, 16.4 ± 0.9 pA. (E) Box plots of dendrite mean GluA2 receptor cluster size pooled from same 3 cultures as shown in (A) and (B). CON, N = 29 dendrites, 0.90 ± 0.06 μm2; TTX, N = 28 dendrites, 1.08 ± 0.10 μm2. (F) Box plots of dendrite mean GluA2 receptor cluster size pooled from the same 5 cultures shown in (C). CON, N = 49 dendrites, 0.83 ± 0.04 μm2; TTX, N = 48 dendrites, 0.98 ± 0.07 μm2. Line series, Student’s paired T test; box plots, Kruskal-Wallis test. Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10 and 90 percentiles; open circles indicate means of individual cells/dendrites; bin size 1 pA, 0.5 μm2; line, median; dot, mean. For all p-values, * with underline indicates significance with p < 0.05.

Revised version: (A) Culture averages of mEPSC amplitudes for untreated (CON) and TTX-treated coverslips (TTX) in each of 3 Rab3A+/+ mouse cortical co-cultures. Culture #1, CON, N = 6, 14.2 ± 2.2 pA; TTX, N = 6, 15.9 ± 1.9 pA; Culture #2, CON, N = 7, 13.8 ± 2.4 pA; TTX, N = 8, 17.7 ± 1.8 pA; Culture #3, CON, N = 10, 13.4 ± 0.8 pA; TTX, N = 9, 15.7 ± 1.0 pA. (B) Culture averages of GluA2 receptor cluster size in the same 3 cultures as shown in (A). Culture #1, CON, N = 10, 0.73 ± 0.09 μm2; TTX, N = 9, 0.89 ± 0.09 μm2; Culture #2, CON, N = 9, 0.91 ± 0.12 μm2; TTX, N = 9, 1.42 ± 0.24 μm2; Culture #3, CON, N = 10, 1.05 ± 0.11 μm2; TTX, N = 9, 0.95 ± 0.15 μm2. (C) Two additional culture averages are included that did not have corresponding mEPSC recordings. Culture #4, CON, N = 10, 0.58 ± 0.07 μm2: TTX, N = 10, 0.71 ± 0.04 μm2; Culture 5, CON, N =10, 0.89 ± 0.05 μm2; TTX, N = 10, 0.95 ± 0.08 μm2. (D) Box plots of cell mean mEPSC amplitudes pooled from same 3 cultures as shown in (A). CON, N = 23 cells, 13.7 ± 0.9 pA; TTX, N = 24 cells, 16.4 ± 0.9 pA. (E) Box plots of dendrite mean GluA2 receptor cluster size pooled from same 3 cultures as shown in (A) and (B). CON, N = 29 dendrites, 0.90 ± 0.06 μm2; TTX, N = 28 dendrites, 1.08 ± 0.10 μm2. (F) Box plots of dendrite mean GluA2 receptor cluster size pooled from the same 5 cultures shown in (C). CON, N = 49 dendrites, 0.83 ± 0.04 μm2; TTX, N = 48 dendrites, 0.98 ± 0.07 μm2. Line series, Student’s paired T test; box plots, Kruskal-Wallis test. Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10 and 90 percentiles; open circles indicate means of individual cells/dendrites; bin size 0.1 pA, 0.05 μm2; line, median; dot, mean. For all p-values, * with underline indicates significance with p < 0.05.

We next pooled data from the multiple cultures to examine effects of activity blockade on GluA2 receptors at the level of cell means. Because the average number of synaptic sites identified in the primary dendrites in three cultures was 20.4, we also took 20 mEPSC samples from each cell, the 21st to the 40th mEPSC recorded, to minimize any differences in variability between GluA2 receptors and mEPSCs that might be due to a difference in sampling. In the data pooled from the 3 matched experiments, neurons from Rab3A+/+ cultures showed a significant increase in mean mEPSC amplitudes following activity blockade (Figure 5D; Table 1). This result indicates that the homeostatic response averaged across the 3 cultures was very similar to, but slightly smaller than, that of the data set presented in Figure 1. The mean for size of GluA2 immunofluorescent receptor clusters showed a trend towards higher values after activity blockade (18.1%) that was of similar magnitude to the mEPSC amplitude increase (19.7%), but did not reach statistical significance for 3 cultures (Figure 5E) or 5 cultures (Figure 5F; Table 1). As expected, mean mEPSC amplitude was not increased following activity blockade in the data pooled from a new set of 3 Rab3A-/- cultures (Table 1). For immunofluorescence results from the same Rab3A-/- cultures, no trend towards higher values was apparent for GluA2 receptor cluster characteristics (Table 1). However, given the variability of the GluA2 response in cultures from Rab3A+/+ mice, it would require a much higher number of Rab3A+/+ and Rab3A-/- cultures to make any firm conclusions about the role of Rab3A in the homeostatic response of GluA2 receptors.

Having used VGLUT1 immunofluorescence to mark synaptic sites, we could also ask whether there was a presynaptic effect of activity blockade on the size, intensity or integral of the VGLUT1 signal at sites apposed to the previously examined GluA2 postsynaptic sites. In the pooled data from cultures prepared from Rab3A+/+ mice, we found no change in the size of the VGLUT1-positive regions, and trends towards reduced intensity and integral (Table 2). These data do not support the idea that activity blockade caused an increase in the amount of VGLUT1 per vesicle, although we cannot rule out that a simultaneous reduction in the number of vesicles obscured the effect, since we did not have an independent label for synaptic vesicles. In cultures prepared from Rab3A-/- mice, the size, intensity and integral of the VGLUT1-positive regions did not significantly change (Table 2).

Earlier version: We have demonstrated here for the first time that the synaptic vesicle protein Rab3A influences the homeostatic regulation of quantal amplitude. What are possible ways that Rab3A could exert its effect? It has previously been shown that exogenous addition of TNFα to hippocampal cultures causes an increase in surface expression of GluA1 receptors (although not GluA2 receptors), and that the homeostatic increase in mEPSC amplitudes is abolished in cultures prepared from the TNFα deletion mouse (Stellwagen et al., 2005; Stellwagen and Malenka, 2006). Furthermore, neurons from TNFα+/+ mice plated on glial feeder layers derived from TNFα-/- mice fail to show the increase in mEPSC amplitude after TTX treatment, indicating that the TNFα inducing the receptor increases following activity blockade comes from the glial cells (Stellwagen and Malenka, 2006). It has recently been shown that the source of TNFα is astrocytes rather than microglia (Heir et al., 2024). Rab3A has been detected in astrocytes (Maienschein et al., 1999; Hong et al., 2016), so to determine whether Rab3A is acting via regulating TNFα release from astrocytes, we performed experiments similar to those of Stellwagen and Malenka (2006) (schema illustrated in Figure 6, Left). Under our culturing conditions, the glial feeder layers would be expected to be composed mainly of astrocytes (Heir et al., 2024). We compared the effect of activity blockade on mEPSC amplitudes recorded from cortical neurons from Rab3A+/+ mice plated on Rab3A+/+ glial feeder layers (Figure 6A); neurons from Rab3A+/+ mice plated on Rab3A-/- glia (Figure 6B), and neurons from Rab3A-/- mice plated on Rab3A+/+ glia (Figure 6C). If Rab3A is required for TNFα release from glia, then Rab3A+/+ neurons plated on Rab3A-/- glia should not show a homeostatic increase in mEPSC amplitude after treatment with TTX, and any cultures with Rab3A+/+ glia should have a normal homeostatic response. We found the opposite result: mEPSC amplitudes increased dramatically in cultures where Rab3A was present in neurons (WT on WT, CON, 13.3 ± 0.5 pA, TTX, 16.7 ± 1.2 pA; WT on KO, CON, 13.3 ± 1.0 pA, TTX, 18.8 ± 1.4 pA), but the increase was greatly diminished in cultures where Rab3A was present only in glia (KO on WT, CON, 15.2 ± 1.1 pA, TTX, 16.9 ± 0.7 pA). In the glial feeder layer culture, the loss of Rab3A in the neurons may have caused an increase in mEPSC amplitude at baseline, similar to the Earlybird mutant in dissociated cultures, but this difference was not significant (CON, WT on WT, 13.3 ± 0.5 pA; CON, KO on WT, 15.2 ± 1.1 pA, p = 0.23, Kruskal-Wallis test), so it is also possible this can be attributed to culture to culture variability. Regarding mEPSC frequency, in cultures with Rab3A+/+ neurons, frequency trended upward after activity blockade, whereas in the culture with Rab3A-/- neurons, frequency was already increased (although not significantly, CON, WT on WT, 2.54 ± 0.57 sec-1; CON, KO on WT, 4.46 ± 1.21 sec-1, p = 0.23, Kruskal-Wallis test), and trended downward after activity blockade.

Revised version: We have demonstrated here for the first time that the synaptic vesicle protein Rab3A influences the homeostatic regulation of quantal amplitude. What are possible ways that Rab3A could exert its effect? It has previously been shown that exogenous addition of TNFα to hippocampal cultures causes an increase in surface expression of GluA1 receptors (although not GluA2 receptors), and that the homeostatic increase in mEPSC amplitudes is abolished in cultures prepared from the TNFα deletion mouse (Stellwagen et al., 2005; Stellwagen and Malenka, 2006). Furthermore, neurons from TNFα+/+ mice plated on glial feeder layers derived from TNFα-/- mice fail to show the increase in mEPSC amplitude after TTX treatment, indicating that the TNFα inducing the receptor increases following activity blockade comes from the glial cells (Stellwagen and Malenka, 2006). It has recently been shown that the source of TNFα is astrocytes rather than microglia (Heir et al., 2024).

Added in the revised version: Rab3A has been detected in astrocytes (Maienschein et al., 1999; Hong et al., 2016), so to determine whether Rab3A is acting via regulating TNFα release from astrocytes, we performed experiments similar to those of Stellwagen and Malenka (2006) (schema illustrated in Figure 6, Left). Under our culturing conditions, the glial feeder layers would be expected to be composed mainly of astrocytes (Heir et al., 2024). We compared the effect of activity blockade on mEPSC amplitudes recorded from cortical neurons from Rab3A+/+ mice plated on Rab3A+/+ glial feeder layers (Figure 6A); neurons from Rab3A+/+ mice plated on Rab3A-/- glia (Figure 6B), and neurons from Rab3A-/- mice plated on Rab3A+/+ glia (Figure 6C). If Rab3A is required for TNFα release from glia, then Rab3A+/+ neurons plated on Rab3A-/- glia should not show a homeostatic increase in mEPSC amplitude after treatment with TTX, and any cultures with Rab3A+/+ glia should have a normal homeostatic response. We found the opposite result: mEPSC amplitudes increased in cultures where Rab3A was present in neurons (WT on WT, CON, 13.3 ± 0.5 pA, TTX, 16.7 ± 1.2 pA; WT on KO, CON, 13.3 ± 1.0 pA, TTX, 18.8 ± 1.4 pA), but the increase was greatly diminished in cultures where Rab3A was present only in glia (KO on WT, CON, 15.2 ± 1.1 pA, TTX, 16.9 ± 0.7 pA). In the glial feeder layer culture, the loss of Rab3A in the neurons may have caused an increase in mEPSC amplitude at baseline, similar to the Earlybird mutant in dissociated cultures, but this difference was not significant (CON, WT on WT, 13.3 ± 0.5 pA; CON, KO on WT, 15.2 ± 1.1 pA, p = 0.78, Tukey’s post-hoc test after two-way ANOVA). Regarding mEPSC frequency, in cultures with Rab3A+/+ neurons, frequency trended upward after activity blockade, whereas in the culture with Rab3A-/- neurons, frequency was already increased (although not significantly, CON, WT on WT, 2.54 ± 0.57 sec-1; CON, KO on WT, 4.46 ± 1.21 sec-1, p = 0.59, Tukey’s), and trended downward after activity blockade.

Rab3A in neurons, not astrocytes, was required for full TTX-induced homeostatic plasticity.

Earlier version: (A)-(C), mEPSC amplitude (Middle) and frequency (Right) data from dissociated cortical neurons plated on an astrocyte feeder layer, each prepared separately from the type of mice depicted in the schema (Left): (A) Neurons from Rab3A+/+ mice plated on astrocytes from Rab3A+/+ mice. Box plots for mEPSC amplitudes (CON, N = 17 cells, 13.3 ± 0.5 pA; TTX, N = 20 cells, 16.7 ± 1.2 pA) and mEPSC frequency in the same recordings (mean, CON, 2.54 ± 0.57 sec-1; TTX, 3.48 ± 0.64 sec-1); from 4 cultures. (B) Neurons from Rab3A+/+ mice plated on astrocytes from Rab3A-/- mice. Box plots for average mEPSC amplitude (CON, N = 11 cells, 13.3 ± 1.0 pA; TTX, N = 11 cells, 18.8 ± 1.4 pA) and average mEPSC frequency in the same recordings (CON, 2.01 ± 0.41 sec-1; TTX, 4.47 ± 1.53 sec-1); from 2 cultures. (C) Neurons from Rab3A-/- neurons plated on astrocytes from Rab3A+/+ mice. Box plots for average mEPSC amplitude (CON, N = 14 cells, 15.2 ± 1.1 pA; TTX, N = 11 cells, 16.9 ± 0.7 pA) and mEPSC frequency in the same recordings (CON, 4.47 ± 1.21 sec-1; TTX, 3.02 ± 0.70 sec-1); from 3 cultures. Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10th and 90th percentiles; open circles represent means from individual cells; line, median; dot, mean. p values (shown on the graphs) are from Kruskal-Wallis test. p- values denoted with * and underline indicates significance of p < 0.05.

Revised version: (A)-(C), mEPSC amplitude (Middle) and frequency (Right) data from dissociated cortical neurons plated on an astrocyte feeder layer, each prepared separately from the type of mice depicted in the schema (Left): (A) Neurons from Rab3A+/+ mice plated on astrocytes from Rab3A+/+ mice. Box plots for mEPSC amplitudes (CON, N = 17 cells, 13.3 ± 0.5 pA; TTX, N = 20 cells, 16.7 ± 1.2 pA) and mEPSC frequency in the same recordings (mean, CON, 2.54 ± 0.57 sec-1; TTX, 3.48 ± 0.64 sec-1); from 4 cultures. (B) Neurons from Rab3A+/+ mice plated on astrocytes from Rab3A-/- mice. Box plots for average mEPSC amplitude (CON, N = 11 cells, 13.3 ± 1.0 pA; TTX, N = 11 cells, 18.8 ± 1.4 pA) and average mEPSC frequency in the same recordings (CON, 2.01 ± 0.41 sec-1; TTX, 4.47 ± 1.53 sec-1); from 2 cultures. (C) Neurons from Rab3A-/- neurons plated on astrocytes from Rab3A+/+ mice. Box plots for average mEPSC amplitude (CON, N = 14 cells, 15.2 ± 1.1 pA; TTX, N = 11 cells, 16.9 ± 0.7 pA) and mEPSC frequency in the same recordings (CON, 4.47 ± 1.21 sec-1; TTX, 3.02 ± 0.70 sec-1); from 3 cultures. Box plot parameters: box ends, 25th and 75th percentiles; whiskers, 10th and 90th percentiles, bin size 0.1 pA; open circles represent means from individual cells; line, median; dot, mean. p values (shown on the graphs) are from Tukey’s post-hoc test following a two-way ANOVA. p-values denoted with * and underline indicates significance of p < 0.05.

In summary, in the absence of Rab3A from glia, the homeostatic plasticity of mEPSC amplitude following activity blockade was normal, whereas in the absence of Rab3A in neurons, homeostatic plasticity was greatly diminished. This result makes it highly unlikely that Rab3A is required for the release of TNFα, or another factor from glia, that induces a homeostatic upregulation of postsynaptic receptors and thereby increases mEPSC amplitude following TTX treatment. Neuronal Rab3A appears to mediate the homeostatic increase in mEPSC amplitude following activity blockade.

Discussion

Earlier version: We found that homeostatic synaptic plasticity of mEPSC amplitude in dissociated mixed cultures of mouse cortical neurons and glia behaved remarkably similar to the mouse NMJ in response to loss of Rab3A function: in cultures from Rab3A-/- mice, the increase in mEPSC amplitudes following prolonged network silencing by TTX was strongly diminished, and in cultures from Rab3AEbd/Ebd mice, basal mEPSC amplitude was increased compared to that of wild-type cultures and was not further modified by TTX treatment. These results suggest that normal function of the synaptic vesicle protein Rab3A is required for the homeostatic increase of mEPSC amplitude in cortical cultures. At the NMJ, the effect of Rab3A on homeostatic upregulation of mEPC amplitude was not mediated by upregulation of postsynaptic receptors. In neurons, the increase in mEPSC amplitude following block of activity was not due to increases in Ca2+-permeable GluA1 homomers because acute NASPM application was unable to reverse the TTX-induced increase in amplitude. Imaging of immunofluorescence of GluA2 receptors revealed the response was more variable than that of mEPSC amplitude. The simplest explanation for our findings at both the NMJ and in neurons is that Rab3A regulates a presynaptic contributor to alteration of quantal amplitude during homeostatic synaptic plasticity.

Revised version: We found that homeostatic synaptic plasticity of mEPSC amplitude in dissociated mixed cultures of mouse cortical neurons and glia behaved remarkably similar to the mouse NMJ in response to loss of Rab3A function: in cultures from Rab3A-/- mice, the increase in mEPSC amplitudes following prolonged network silencing by TTX was strongly diminished, and in cultures from Rab3AEbd/Ebd mice, basal mEPSC amplitude was increased compared to that of wild-type cultures and was not further modified by TTX treatment. These results suggest that normal function of the synaptic vesicle protein Rab3A is required for the homeostatic increase of mEPSC amplitude in cortical cultures.

Added in the revised version: Is Rab3A acting in the presynaptic cell to cause homeostatic increase in mEPSC amplitude?

Added in the revised version: The impetus for our current study was two previous studies in which we examined homeostatic regulation of quantal amplitude at the NMJ. An advantage of studying the NMJ is that synaptic ACh receptors are easily identified with fluorescently labeled α-bungarotoxin, which allows for very accurate quantification of postsynaptic receptor density. We were able to detect a known change due to mixing 2 colors of α-bungarotoxin to within 1% (Wang et al., 2005). Using this model synapse, we showed that there was no increase in synaptic AChRs after TTX treatment, whereas miniature endplate current increased 35% (Wang et al., 2005). We further showed that the presynaptic protein Rab3A was necessary for full upregulation of mEPC amplitude (Wang et al., 2011). These data strongly suggested Rab3A contributed to homeostatic upregulation of quantal amplitude via a presynaptic mechanism. With the current study showing that Rab3A is required for the homeostatic increase in mEPSC amplitude in cortical neuron cultures, one interpretation is that in both situations, Rab3A is required for an increase in the presynaptic quantum.

Added in the revised version: The presynaptic quantum, or the amount of transmitter released during a single fusion event, is an important contributor to quantal size. The amount of transmitter released during vesicle fusion can be affected by the kinetics of the fusion pore opening (Chang et al., 2017). We previously used amperometric measurements in mouse adrenal chromaffin cells to show loss of Rab3A increased the occurrence of very small amplitude fusion pore feet (Wang et al., 2008). Paired with our finding that there is an increase in the occurrence of very slow-rising, abnormally shaped mEPCs at the neuromuscular junction of Rab3A-/- mice, our data suggest that small synaptic vesicles may have a fusion pore step under certain circumstances (see also (Chiang et al., 2018)). If activity blockade causes an increase in mEPSC amplitude due to a more rapid opening of a fusion pore, or a larger fusion pore conductance, it is possible that Rab3A is required for this modulation.

Added in the revised version: Another way to increase the presynaptic quantum is to increase levels of the transmitter transporter (Daniels et al., 2004; De Gois et al., 2005; Wilson et al., 2005; Hartman et al., 2006). However, we found that immunofluorescence for the glutamate transporter, VGLUT1, was not increased after activity blockade under our experimental conditions. A similar lack of increase in VGLUT1 levels, relative to synapsin, was previously observed in hippocampal cultures treated with NBQX, an AMPA receptor blocker (see supplemental data, (Wilson et al., 2005)). Taken together with our findings, it appears that there is not strong evidence that the amount of VGLUT1 at synapses is a contributing factor to the increase in mEPSC amplitude in neuronal cultures after activity blockade with TTX.

Added in the revised version: Finally, the amount of transmitter released by a vesicle could be increased due to larger vesicle size. Larger vesicle diameter has been observed following activity disruption in multiple paradigms, including ocular deprivation, cortical ablation, and immobilization (Kawana et al., 1971; Lloret and Saavedra, 1975; Cheresharov et al., 1978). Loss of Rab3 family members is associated with increased vesicle size (Riedel et al., 2002; Schonn et al., 2010). An increase in volume has been shown to increase miniature endplate junctional current amplitude at the Drosophila NMJ (Karunanithi et al., 2002). Rab3A could be required for an activity-dependent increase in vesicle size, but the increase in diameter needed for a 25% increase in transmitter, and therefore mEPSC amplitude (∼3%), would be difficult to detect.

Added in the revised version: Is Rab3A acting on postsynaptic receptors to cause increase in mEPSC amplitude?

Added in the revised version: In neurons, the increase in mEPSC amplitude following block of activity was not due to increases in Ca2+-permeable GluA1 homomers because acute NASPM application was unable to reverse the TTX-induced increase in amplitude. While we found that the effect of TTX treatment on mEPSC amplitude and GluA2 receptor levels were of similar magnitude, the effect on GluA2 receptor levels was not statistically significant for the pooled data sets acquired in the same 3 cultures, nor in a larger data set containing an additional 2 cultures. In light of the extensive literature documenting increases in GluA1 or GluA2 receptors, or both, after activity blockade, the less than convincing increase in GluA2 receptor levels in the current work came as a surprise. The reason for the lack of significance was that the increase in receptors was more variable. We are aware of three possible explanations for greater variability. 1) The number of synaptic sites sampled was smaller for the imaging data set. This was because we limited our analysis to synaptic receptors on the primary dendrite, where a clear increase in sensitivity to exogenously applied glutamate was demonstrated (see Figure 3 in (Turrigiano et al., 1998)). To address this issue, we limited our sample size of mEPSCs in the matched experiments to the same number per cell as the mean sample size per cell for imaging, but there was still a discrepancy in the effect of TTX on mEPSCs vs. receptors. 2) The small sample of only 3 cultures was not enough to reveal the increase in receptors. We attempted to address this by adding 2 more cultures. Although both additional cultures showed small increases in GluA2 receptor cluster size, when cells from these cultures were included in the pooled data set, the p value was not substantially improved. A remaining possibility is: 3) The complexity involved in processing cultures for immunofluorescence introduces variability that obscures the true effect. One can get around technical issues related to immunofluorescence processing by performing Western blots, but then any changes detected could be due to upregulation of extrasynaptic receptors.

Added in the revised version: What remains unclear is why our data were impacted by the variability inherent in immunofluorescence measurements, but data of other studies were not. Some other studies have used the number of synaptic sites or puncta as the sample size (for example, (Hou et al., 2008; Ibata et al., 2008; Wang et al., 2019)). When we analyzed our data after including as few as 6 randomly selected cluster size measurements per cell, a Kruskal-Wallis test gave highly significant p values for 3 cultures (p = 0.001, n = 174 sites (6 *29 cells)) and 5 cultures (p = 0.005, n = 294 sites (6 *49 cells)). We found only a single study that had comparable sample sizes for receptors and mEPSC amplitudes, and the p value was < 0.01 for mEPSC amplitudes (n = 9 and 10 cells for CON and TTX, respectively), but five times higher (< 0.05) for receptor data (n = 13 and 14 cells for CON and TTX) (Xu and Pozzo-Miller, 2017). This single example is consistent with our suggestion that measurements of changes in receptors after activity blockade are more variable than measurements of mEPSC amplitudes. In sum, the lack of a robust increase in receptor levels leaves open the possibility that there is a presynaptic contributor to quantal size in mouse cortical cultures. However, the variability could arise from technical factors associated with the immunofluorescence method, and the mechanism of Rab3A-dependent homeostatic plasticity could be presynaptic for the NMJ and postsynaptic for cortical neurons.

Added in the revised version: Increased receptors likely contribute to increases in mEPSC amplitudes, but because we do not have a significant increase in GluA2 receptors in our experiments, it is impossible to conclude that the increase is lacking in cultures from Rab3A-/- neurons. Furthermore, without additional experiments targeting Rab3A deletion to the pre-or post-synaptic compartment, we cannot rule out that Rab3A is regulating homeostatic plasticity in mouse cortical cultures by acting postsynaptically on receptor trafficking. To our knowledge, there is no evidence localizing Rab3A to postsynaptic dendritic sites, although other presynaptic molecules, such as SNARE proteins, synaptotagmins, and NSF, have been identified in the postsynaptic compartment (Lledo et al., 1998; Nishimune et al., 1998; Osten et al., 1998; Song et al., 1998; Araki et al., 2010; Kennedy et al., 2010; Suh et al., 2010; Jurado et al., 2013; Hussain and Davanger, 2015; Gu et al., 2016; Wu et al., 2017). A recent report implicates Rab3A in the localization of plasma membrane proteins, including the EGF receptor, in rafts in HEK cells and Jurkat-T cells, which indicates Rab3A can be found in other specialized locations besides the presynaptic nerve terminal (Diaz-Rohrer et al., 2023).

Added in the revised version: Neuronal, not glial, Rab3A is required for homeostatic regulation of quantal amplitude

Added in the revised version: Astrocytic release of TNFα was shown to mediate the increase in mEPSC amplitudes in prolonged TTX-treated cultures of dissociated hippocampal neurons and cultures of hippocampal slices (Stellwagen and Malenka, 2006; Heir et al., 2024). TNF-α accumulates after activity blockade, and directly applied to neuronal cultures, can cause an increase in GluA1 receptors, providing a potential mechanism by which activity blockade leads to the homeostatic upregulation of postsynaptic receptors (Beattie et al., 2002; Stellwagen et al., 2005; Stellwagen and Malenka, 2006). However, we found that loss of Rab3A in glia did not disrupt the increase in mEPSC amplitude after activity blockade, so it is unlikely that the Rab3A-dependent homeostatic increase in mEPSC amplitude is via the astrocytic-TNF-α pathway. Interestingly, we previously showed that the homeostatic increase in NMJ mEPC amplitude was completely normal in the absence of TNF-α (Wang et al., 2011), suggesting the possibility that neuronal Rab3A can act via a non-TNF-α mechanism to contribute to homeostatic regulation of quantal amplitude, although we have not ruled out a neuronal Rab3A-mediated TNF-α pathway in cortical cultures.

Is the lack of homeostatic plasticity in the Earlybird mutant due to occlusion?

We found that mEPSC amplitude was increased in untreated cultures prepared from Rab3AEbd/Ebd mice, compared to mEPSC amplitude in cultures from wild type mice. If there is a limit to how large an mEPSC can become, it is possible that the mutant Rab3A does not affect homeostatic plasticity, only occludes it. Similar increases in mEPSC amplitude at baseline, combined with a failure to increase further after activity blockade, were noted in cultures prepared from mECP2, AKAP, Homer1a, and Arc/Arg3.1 deletion mice (Shepherd et al., 2006; Hu et al., 2010; Xu and Pozzo-Miller, 2017; Sanderson et al., 2018) as well as in cultures treated with ubiquitin proteasome inhibitors (Jakawich et al., 2010) or microRNA 186-5p (Silva et al., 2019), suggesting the possibility of a generalized phenomenon in which increased mEPSC amplitude is induced when homeostatic regulation is disrupted. Since we did not observe such an increase in mEPSC amplitude at baseline in cultures from Rab3A-/- mice, it remains a possibility that the point mutant may bind to novel partners, causing activities that would not be either facilitated or inhibited by Rab3A. Still, it is a strong coincidence that the novel activity of the mutant Earlybird Rab3A would affect mEPSC amplitude, the same characteristic that is modulated by activity blockade in a Rab3A dependent manner. Although we cannot rule out occlusion, an alternate interpretation is that the presence of the mutant protein mimics the condition of activity blockade.

Added in the revised version: Rab3A is the first presynaptic function protein required for full homeostatic plasticity of mEPSC amplitude

Removed in the revised version: Homeostatic plasticity of postsynaptic receptors was less consistent than that of mEPSC amplitude

Removed in the revised version: While we found that the effect of TTX treatment on mEPSC amplitude and GluA2 receptor levels were of similar magnitude, the effect on GluA2 receptor levels was not statistically significant for the pooled data sets acquired in the same cultures, nor in a larger data set containing 5 cultures. In light of the extensive literature documenting increases in GluA1 or GluA2 receptors, or both, after activity blockade, the less than convincing increase in GluA2 receptor levels in the current work came as a surprise. The reason for the lack of significance was that the increase in receptors was more variable. We are aware of three possible explanations for greater variability. 1) The number of synaptic sites sampled was smaller for the imaging data set. This was because we limited our analysis to synaptic receptors on the primary dendrite, where the increase in sensitivity to exogenously applied glutamate was shown to be greatest (see Figure 3 in (Turrigiano et al., 1998)). To address this issue, we limited our sample size of mEPSCs in the matched experiments to the same number per cell as the mean sample size per cell for imaging, but there was still a discrepancy in the effect of TTX on mEPSCs vs. receptors. 2) The small sample of only 3 cultures was not enough to reveal the increase in receptors. We attempted to address this by adding 2 more cultures. Although both additional cultures showed increases in GluA2 receptor cluster size, the magnitudes of the increases were not large, and when cells from these cultures were included in the pooled data set, the p value was not substantially improved. The only remaining possibility is: 3) The complexity involved in processing cultures for immunofluorescence introduces variability that obscures the true effect. One can get around technical issues related to immunofluorescence processing by performing Western blots, but then any changes detected could be due to upregulation of extrasynaptic receptors. What remains unclear is why our data were impacted by this but data of other studies were not. Some other studies have used the number of synaptic sites or puncta as the sample size (for example, (Hou et al., 2008; Ibata et al., 2008; Wang et al., 2019)). When we analyzed our data after including as few as 6 randomly selected cluster size measurements per cell, a Kruskal-Wallis test gave highly significant p values for 3 cultures (p = 0.001, n = 174 sites (6 *29 cells)) and 5 cultures (p = 0.005, n = 294 sites (6 *49 cells)). We found only a single study that had comparable sample sizes for receptors and mEPSC amplitudes, and the p value was < 0.01 for mEPSC amplitudes (n = 9 and 10 cells for CON and TTX, respectively), but five times higher (< 0.05) for receptor data (n = 13 and 14 cells for CON and TTX) (Xu and Pozzo-Miller, 2017). This single example is consistent with our suggestion that measurements of changes in receptors after activity blockade are more variable than measurements of mEPSC amplitudes. In sum, our data suggest that TTX treatment of mouse cortical cultures consistently leads to an increase in mEPSC amplitudes, and sometimes leads to an increase in synaptic GluA2 receptor cluster size. Increased receptors likely contribute to increases in mEPSC amplitudes, but our current data, coupled with our previous findings at the mouse neuromuscular junction, support the idea that there are additional sources contributing to the homeostatic increase in quantal size.

Removed in the revised version: Neuronal, not glial, Rab3A is required for homeostatic regulation of quantal amplitude

Removed in the revised version: Astrocytic release of TNFα was shown to mediate the increase in mEPSC amplitudes in prolonged TTX-treated cultures of dissociated hippocampal neurons and cultures of hippocampal slices (Stellwagen and Malenka, 2006; Heir et al., 2024), but we found that loss of Rab3A in glia did not disrupt the increase in mEPSC amplitude after activity blockade. Further, we previously showed that the homeostatic increase in NMJ mEPC amplitude was completely normal in the absence of TNFα (Wang et al., 2011). This leaves neuronal Rab3A as the potential controller of homeostatic regulation of quantal amplitude.

Removed in the revised version: Is Rab3A acting in the presynaptic cell to cause homeostatic increase in mEPSC amplitude?

Removed in the revised version: Our current results support the surprising possibility that the presynaptic vesicle protein Rab3A either regulates postsynaptic AMPARs or regulates a presynaptic contributor to homeostatic synaptic plasticity. To our knowledge, there is no evidence localizing Rab3A to postsynaptic dendritic sites. Other presynaptic molecules, such as SNARE proteins, synaptotagmins, and NSF, have been identified in the postsynaptic compartment (Lledo et al., 1998; Nishimune et al., 1998; Osten et al., 1998; Song et al., 1998; Araki et al., 2010; Kennedy et al., 2010; Suh et al., 2010; Jurado et al., 2013; Hussain and Davanger, 2015; Gu et al., 2016; Wu et al., 2017). A recent report implicates Rab3A in the localization of plasma membrane proteins, including the EGF receptor, in rafts in HEK cells and Jurkat-T cells, which indicates Rab3A can be found in other specialized locations besides the presynaptic nerve terminal (Diaz-Rohrer et al., 2023). Without additional experiments targeting Rab3A deletion to the pre- or post-synaptic compartment, we cannot rule out that Rab3A is regulating homeostatic plasticity in mouse cortical cultures by acting postsynaptically on receptor trafficking. However, given our previous finding of a lack of any increase in AChR after TTX treatment the NMJ (Wang et al., 2011), the simplest interpretation is that in these two situations, Rab3A is required for an increase in the presynaptic quantum rather than an increase in postsynaptic receptors.

Removed in the revised version: The presynaptic quantum, or the amount of transmitter released during a single fusion event, is an important contributor to quantal size. The amount of transmitter released during vesicle fusion can be affected by the kinetics of the fusion pore opening (Chang et al., 2017). We previously used amperometric measurements in mouse adrenal chromaffin cells to show loss of Rab3A increased the occurrence of very small amplitude fusion pore feet (Wang et al., 2008). Paired with our finding that there is an increase in the occurrence of very slow-rising, abnormally shaped mEPCs at the neuromuscular junction of Rab3A-/- mice, our data suggest that small synaptic vesicles may have a fusion pore step under certain circumstances (see also (Chiang et al., 2018)). If activity blockade causes an increase in mEPSC amplitude due to a more rapid opening of a fusion pore, or a larger fusion pore conductance, it is possible that Rab3A is required for this modulation.

Removed in the revised version: Another way to increase the presynaptic quantum is to increase levels of the transmitter transporter (Daniels et al., 2004; De Gois et al., 2005; Wilson et al., 2005; Hartman et al., 2006). However, we found that immunofluorescence for the glutamate transporter, VGLUT1, was not increased after activity blockade under our experimental conditions. A similar lack of increase in VGLUT1 levels, relative to synapsin, was previously observed in hippocampal cultures treated with NBQX, an AMPA receptor blocker (see supplemental data, (Wilson et al., 2005)). Taken together with our findings, it appears that there is not strong evidence that the amount of VGLUT1 at synapses is a contributing factor to the increase in mEPSC amplitude in neuronal cultures after activity blockade with TTX.

Removed in the revised version: Finally, the amount of transmitter released by a vesicle could be increased due to larger vesicle size. Larger vesicle diameter has been observed following activity disruption in multiple paradigms, including ocular deprivation, cortical ablation, and immobilization (Kawana et al., 1971; Lloret and Saavedra, 1975; Cheresharov et al., 1978). Loss of Rab3 family members is associated with increased vesicle size (Riedel et al., 2002; Schonn et al., 2010). Finally, an increase in volume has been shown to increase miniature endplate junctional current amplitude at the Drosophila NMJ (Karunanithi et al., 2002). Rab3A could be required for an activity-dependent increase in vesicle size, but the increase in diameter needed for a 25% increase in transmitter, and therefore mEPSC amplitude, would be difficult to detect.

The number of molecules shown to be required for homeostatic synaptic plasticity of quantal size in neuronal cultures may have surpassed ∼20 (see (Koesters et al., 2024) for a recent summary), but fall into two main categories: molecules involved in glutamate receptor expression or trafficking, and signaling molecules. To our knowledge, this is the first evidence that a presynaptic protein is participating in this process. Homeostatic synaptic plasticity is becoming increasingly implicated in essential normal functions such as sleep (Tononi and Cirelli, 2014; Diering et al., 2017; Torrado Pacheco et al., 2021), and pathological conditions such as epilepsy, neuropsychiatric disorders, Huntington’s, and alcohol use disorder (Trasande and Ramirez, 2007; Fernandes and Carvalho, 2016; Wang et al., 2017; Lovinger and Abrahao, 2018; Smith-Dijak et al., 2019; Lignani et al., 2020; Suzuki et al., 2021; Kavalali and Monteggia, 2023). Therefore, our identification of a novel Rab3A-dependent regulatory pathway in homeostatic synaptic plasticity may have important therapeutic implications. Our findings also further cement the importance of Rab3A in activity-dependent modulation of synaptic strength.

Materials and Methods

Animals

To determine the role of Rab3A in homeostatic plasticity in mouse cortical cultures, we employed two distinct genetic mouse strains with altered Rab3A function. Rab3A deletion mice were generated as follows: Rab3A+/- heterozygous mice were bred and genotyped as previously described (Kapfhamer et al., 2002; Wang et al., 2008), and maintained as a colony using multiple heterozygous breeding pairs. Homozygous Rab3A+/+ or Rab3A-/- mouse pups were obtained by breeding Rab3A+/+ pairs and Rab3A-/- pairs, respectively, so that the pups were the first-generation progeny of a homozygous mating. This protocol minimizes the tendency of multiple generations of Rab3A+/+ and Rab3A-/- homozygous breedings to produce progeny that are more and more genetically distinct. Rab3AEbd/Ebdmice were identified in an EU-mutagenesis screen of C57BL/6J mice, and after a cross to C3H/HeJ, were backcrossed for 3 generations to C57BL/6J (Kapfhamer et al., 2002). Rab3A+/Ebd heterozygous mice were bred at Wright State University and genotyped in a two-step procedure: 1. a PCR reaction (forward primer: TGA CTC CTT CAC TCC AGC CT; reverse primer: TGC ACT GCA TTA AAT GAC TCC T) followed by 2. a restriction digest with enzyme Bsp1286I (New England Biolabs) that distinguishes the Earlybird mutant by its different base-pair products. Rab3A+/- mice were backcrossed with Rab3A+/+ mice from the Earlybird heterozygous colony for 11 generations in an attempt to establish a single wild type strain, but differences in mEPSC amplitude and adrenal chromaffin cell calcium currents persisted, likely due to genes that are close to the Rab3A site, resulting in two wild type strains: 1. Rab3A+/+ from the Rab3A+/- colony, and 2. Rab3A+/+ from the Rab3A+/Ebd colony.

Primary Culture of Mouse Cortical Neurons

Primary dissociated cultures of mixed neuronal and glia populations were prepared as previously described (Hanes et al., 2020). Briefly, postnatal day 0-2 (P0-P2) Rab3A+/+, Rab3A-/- or Rab3AEbd/Ebd neonates were euthanized by rapid decapitation, as approved by the Wright State University Institutional Animal Care and Use Committee, and brains were quickly removed. Each culture was prepared from the cortices harvested from two animals; neonates were not sexed. Cortices were collected in chilled Neurobasal-A media (Gibco) with osmolarity adjusted to 270 mOsm and supplemented with 40 U/ml DNAse I (ThermoFisher Scientific). The tissues were digested with papain (Worthington Biochemical) at 20 U/ml at 37°C for 20 minutes followed by trituration with a sterile, fire-polished Pasteur pipette, then filtered through a 100 μm cell strainer, and centrifuged at 1100 rpm for 2 minutes. After discarding the supernatant, the pellet was resuspended in room temperature Neurobasal-A media (270 mOsm), supplemented with 5% fetal bovine serum for glia growth, and 2% B-27 supplement to promote neuronal growth (Gibco), L-glutamine, and gentamicin (ThermoFisher Scientific). Neurons were counted and plated at 0.15 * 106 cells/coverslip onto 12 mm coverslips pre-coated with poly-L-lysine (BioCoat, Corning). This plating density produces a “high density” culture characterized by a complex mesh of neuronal processes criss-crossing the field of view, completely filling the space between cell bodies (see Figure 4). We found in preliminary experiments that an increase in mEPSC amplitude in mouse cortical cultures was inconsistent when cell density was low, likely due to lower levels of baseline activity, although it is a limitation of this study that we did not directly measure activity levels in cultures prepared from Rab3A+/+, Rab3A-/- or Rab3AEbd/Ebd mice, and therefore do not know if differences in results are due to differences in activity levels. While unlikely for Rab3A+/+ and Rab3A-/- studies, given the mEPSC amplitude in untreated cultures was relatively small, it remains possible that in Rab3AEbd/Ebd studies, the larger mEPSC amplitude in untreated cultures was due to a loss of activity in these cultures. The culture media for the first day (0 DIV) was the same as the above Neurobasal-A media supplemented with FBS, B-27, L-glutamine, and gentamicin, and was switched after 24 hours (1 DIV) to media consisting of Neurobasal-A (270 mOsm), 2% B-27, and L-glutamine without FBS to avoid its toxic effects on neuronal viability and health (Stellwagen and Malenka, 2006). Half of the media was changed twice weekly and experiments were performed at 13-14 DIV. Two days prior to experiments, tetrodotoxin (TTX) (500 nM; Tocris), a potent Na+ channel blocker, was added to some cultures to chronically silence all network activity and induce homeostatic synaptic plasticity mechanisms, while untreated sister cultures served as controls. Cultures prepared from mutant mice were compared with cultures from wild-type mice from their respective colonies. Note that the cultures comprising the Rab3A+/+ data here are a subset of the data previously published in Hanes et al., 2020. This smaller data set was restricted to the time period over which cultures were also prepared from Rab3A-/- mice.

Preparation of Glial Feeder Layers

Glial feeder layers were prepared from the cortices of P0-P2 Rab3A+/+ or Rab3A-/- mouse pups as described previously (Stellwagen and Malenka, 2006). Briefly, cortices were dissected and cells were dissociated as described above. Cell suspensions of mixed neuronal and glial populations were plated onto glass coverslips pre-coated with poly-L-lysine in Dulbecco’s Modified Eagle Media (ThermoFisher Scientific) supplemented with 5% FBS (to promote glial proliferation and to kill neurons), L-glutamine, and gentamicin, and maintained in an incubator at 37°C, 5% CO2; cultures were maintained in this manner for up to 1 month to generate purely glial cultures (all neurons typically died off by 7 DIV). Culture media was replaced after 24 hours, and subsequent media changes were made twice weekly, replacing half of the culture media with fresh media. Feeder layers were not used for neuronal seeding until all native neurons were gone and glia approached 100% confluency (visually inspected).

Plating of Neurons on Glial Feeder Layers

Cortical neurons were obtained as described above. The cell pellet obtained was resuspended in Neurobasal-A (osmolarity adjusted to 270 mOsm) containing B27 (2%, to promote neuronal growth), L-glutamine, and 5-fluorodeoxyuridine (FdU, a mitotic inhibitor; Sigma). Addition of FdU was used to prevent glial proliferation and contamination of the feeder layer with new glia, promoting only neuronal growth on the feeder layers (FdU-containing media was used for the maintenance of these cultures and all subsequent media changes). Glial culture media was removed from the feeder layer cultures, and the neuronal cell suspension was plated onto the glial feeder cultures. The culture strategy used to distinguish the relative roles of neuronal and glial Rab3A is outlined in Figure 8, Left. At 1 DIV, all of the culture media was removed and replaced with fresh Neurobasal-A media containing FdU described above, and half of the media was replaced twice per week for all subsequent media changes. Cultures were maintained in a 37°C, 5% CO2 incubator for 13-14 DIV.

Whole-Cell Voltage Clamp to Record mEPSCs

At 13-14 DIV, mEPSCs from TTX-treated and untreated sister cultures of Rab3A+/+ or Rab3A-/- neurons from the Rab3A+/- colony, or Rab3A+/+ or Rab3AEbd/Ebd neurons from the Rab3A+/Ebd colony, were recorded via whole-cell voltage clamp to assess the role of Rab3A in homeostatic synaptic plasticity. Recordings were taken from pyramidal neurons, which were identified visually by a prominent apical dendrite; images were taken of all cells recorded from. Cells were continuously perfused with a solution consisting of (in mM): NaCl (115), KCl (5), CaCl2 (2.5), MgCl2 (1.3), dextrose (23), sucrose (26), HEPES (4.2), pH = 7.2 (Stellwagen and Malenka, 2006). On the day of recording, the osmolarity of the media from the cultures was measured (normally 285 – 295 mOsm) and the perfusate osmolarity was adjusted to match the culture osmolarity to protect against osmotic shock to the neurons. To isolate glutamatergic mEPSCs, TTX (500 nM) and picrotoxin (50 μM), were included in the perfusion solution to block action potentials and GABAergic currents, respectively. The NMDA receptor antagonist, APV, was not included in the perfusion solution because in pilot studies, all mEPSCs were blocked by 10 µM CNQX and 50 μm picrotoxin, demonstrating no APV-sensitive mEPSCs were present (data not shown). Previous studies have described NMDA mEPSCs in neuronal cultures, but recordings were performed in 0 Mg2+ (Watt et al., 2000; Sutton et al., 2006). In our recordings, the presence of extracellular Mg2+ (1.3 mM) and TTX would mean that the majority of NMDA receptors are blocked by extracellular Mg2+ at resting Vm. Patch electrodes (3.5 – 5 MΩ) were filled with an internal solution containing (in mM): K-gluconate (130), NaCl (10), EGTA (1), CaCl2 (0.132), MgCl2 (2), HEPES (10), pH = 7.2. Osmolarity was adjusted to 10 mOsm less than the perfusion solution osmolarity, which we have found improves success rate of achieving low access resistance recordings. Neurons were clamped at a voltage of -60 mV using an Axopatch 200B patch-clamp (Axon Instruments), recorded from for 2 – 5 minutes, and data were collected with Clampex 10.0/10.2 (Axon Instruments).

NASPM application

The antagonist of Ca2+-permeable AMPA receptors, N-naphthyl acetylspermine (NASPM, 20 μM; Tocris), was applied during recordings in a subset of experiments. NASPM is a synthetic analog of Joro Spider Toxin (JSTX) (Koike et al., 1997), and chemicals in this family block Ca2+ permeable glutamate receptors (Blaschke et al., 1993; Herlitze et al., 1993; Iino et al., 1996). The presence of the edited GluA2 subunit disrupts both the Ca2+ permeability and the sensitivity to JSTX and related substances (Blaschke et al., 1993; Bochet et al., 1994; Jonas and Burnashev, 1995). Therefore, 20 μM NASPM is expected to completely block AMPA receptors containing GluA1, GluA3 or GluA4 subunits (Hollmann et al., 1991; Burnashev et al., 1992), or the Kainate receptors containing GluA5 or GluA6 (Koike et al., 1997; Sun et al., 2009), but be ineffective against any heteromer or homomer containing GluA2 subunits. Because the effect of the spider toxins and their analogues are use-dependent (Herlitze et al., 1993; Iino et al., 1996; Koike et al., 1997), NASPM was applied with a depolarizing high K+ solution (25 mM KCl, 95 mM NaCl) which we expected to trigger release of glutamate and opening of glutamate-activated receptors, allowing entry of NASPM into the pore. Baseline recordings were performed for 2 minutes in our standard perfusate, followed by perfusion of NASPM + High K+ solution for 45 seconds, and then perfusion of a NASPM only solution for 5 minutes, after which recording was recommenced for 5 minutes (because we found in pilot experiments that frequency was reduced following NASPM application).

Analysis of mEPSCs

Miniature excitatory postsynaptic currents (mEPSCs) were manually selected using Mini Analysis software (Synaptosoft), a standard program used for mEPSC event detection and analysis. Records were filtered at 2 kHz using a low-pass Butterworth filter prior to selection. The program detection threshold was set at 3 pA but the smallest mEPSC selected was 4.01 pA. A fully manual detection process means that the same criterion (“does this look like an mEPSC?”) is applied to all events, not just the false positives or the false negatives, which prevents the bias from being primarily at one end or the other of the range of mEPSC amplitudes. It is important to note that when performing the MiniAnalysis process, the researcher did not know whether a record was from an untreated cell or a TTX-treated cell.

Immunofluorescence, microscopy, and data analysis

Primary cultures of mouse cortical neurons were grown for 13-14 DIV. Antibodies to GluA2 (mouse ab against N-terminal, EMD Millipore) were added directly to live cultures at 1:40 dilution, and incubated at 37 °C in a CO2 incubator for 45 minutes. Cultures were rinsed 3 times with PBS/5% donkey serum before being fixed with 4% paraformaldehyde. After 3 rinses in PBS/5% donkey serum, cultures were incubated in CY3-labeled donkey-anti-mouse secondary antibodies for 1 hour at room temperature, rinsed in PBS/5% donkey serum, permeabilized with 0.2% saponin, and incubated in chick anti-MAP2 (1:2500, AbCAM) and rabbit anti-VGLUT1 (1:4000, Synaptic Systems) for 1 hour at room temperature in PBS/5% donkey serum. After rinsing with PBS/5% donkey serum, coverslips were incubated with 488-anti chick and CY5-anti rabbit secondary antibodies for 1 hour at room temperature, rinsed, blotted to remove excess liquid, and inverted on a drop of Vectashield (Vector Labs). Coverslips were sealed with nail polish and stored at 4 °C for < 1 week before imaging. All secondary antibodies were from Jackson Immunoresearch and were used at 1:225 dilution.

Coverslips were viewed on a Fluoview FV1000 laser scanning confocal microscope with a 60x oil immersion, 1.35 NA objective. Once a pyramidal neuron was identified and a single confocal section of the cell acquired (see Top, Figure 4), Fluoview 2.1 software was used to zoom in on the primary dendrite (5X) and a series of confocal sections were taken every 0.5 μm. Images were analyzed offline with ImagePro 6 (Cybernetics). The composite image was used to locate synaptic sites containing both VGLUT1 and GluA2 immunoreactivity in close apposition to each other and to the primary dendrite or a secondary branch identified with MAP2 immunoreactivity. An area of interest (AOI) was manually drawn around the GluA2 cluster in the confocal section in which it was the brightest. The AOIs for a dendrite were saved in a single file; the AOI number and the confocal section it was associated with were noted for later retrieval. For quantification, AOIs were loaded, an individual AOI was called up on the appropriate section, and the count/measurement tool used to apply a threshold. Multiple thresholds were used, 400-550, but we found that within a given experiment, the magnitude of the TTX effect was not altered by different thresholds, therefore, all data presented here use a threshold of 450. Pixels within the cluster that were above the threshold were automatically outlined, and size, average intensity, and integral of the outlined region saved to Origin2020 (OriginLab) file for calculating mean values. For analyzing VGLUT1 sites, previous AOIs were examined, and if not directly over the VGLUT1 site, moved to encompass it (this happened when GluA2 and VGLUT1 were side by side rather than overlapping). The threshold was 200 for all VGLUT1 data except Rab3A+/+ Culture #1, where threshold was 400. There was no threshold that worked for both that culture and the other cultures, with too much of the background regions being included for Culture #1 with a threshold of 200, and too little of the synaptic site area being included for the other Cultures with a threshold of 400.

Statistics

Earlier version: In our previous publication (Hanes et al., 2020), we described the characteristics of homeostatic plasticity of mEPSC amplitudes in mouse cortical cultures using multiple mEPSC amplitude quantiles from each cell to create data sets with thousands of samples. We compared the cumulative distribution functions for mEPSC amplitudes in untreated vs. TTX- treated cultures, as well as sorting the values from smallest to largest, and plotted TTX values vs. control values in the rank order plot first described in (Turrigiano et al., 1998) to demonstrate synaptic scaling. Finally, we plotted the ratio of TTX values/control values vs. control values to show that the ratio, which represents the effect of TTX, increased with increasing control values, which we termed “divergent scaling.” Here, we wanted to avoid the inflation of sample size caused by pooling multiple mEPSCs per cell in order to statistically compare the effects of TTX in cultures of wild type mice, Rab3A-/- mice and Rab3AEbd/Ebd mutant mice. Therefore, we computed the mean mEPSC amplitude for each cell and compared means from untreated and TTX-treated cultures with the non-parametric Kruskal-Wallis test, with N = the number of cells. The overall means are presented ± SEM in the text, and p values are displayed above each plot. The same process was carried out for the characteristics of synaptic GluA2 receptor clusters and VGLUT1 sites in imaging experiments, with N = the number of dendrites (one dendrite was sampled for each cell). Although we quantified three characteristics of the GluA2 clusters and VGLUT1 sites, size, average intensity, and integral, we have not corrected p values for multiple comparisons. Means are presented as box plots + data, with the box at 25 and 75 percentiles, whiskers at 10 and 90 percentiles, a dot indicating the mean, and a line indicating the median. In Figure 3, mEPSC amplitudes were measured within the same cell before and after NASPM application, so a paired t test was used to compare the pre- and post-NASPM mEPSC amplitudes. In Figure 5A-C the mean mEPSC amplitude for untreated and TTX-treated coverslips in the same culture are compared with a paired t test. Statistics were computed and plots were created using Origin2020 (OriginLab).

Revised version: In our previous publication (Hanes et al., 2020), we described the characteristics of homeostatic plasticity of mEPSC amplitudes in mouse cortical cultures using multiple mEPSC amplitude quantiles from each cell to create data sets with thousands of samples. We compared the cumulative distribution functions for mEPSC amplitudes in untreated vs. TTX-treated cultures, as well as sorting the values from smallest to largest, and plotted TTX values vs. control values in the rank order plot first described in (Turrigiano et al., 1998) to demonstrate synaptic scaling. Finally, we plotted the ratio of TTX values/control values vs. control values to show that the ratio, which represents the effect of TTX, increased with increasing control values, which we termed “divergent scaling.” Here, we wanted to avoid the inflation of sample size caused by pooling multiple mEPSCs per cell. To statistically compare the effects of TTX in cultures of wild type mice, Rab3A-/- mice and Rab3AEbd/Ebd mutant mice, or, WT-WT, WT-KO and KO-WT neuron-glial cultures, we performed a two-way ANOVA, followed by a Tukey test, and used the p value for each specific situation (for example, CON vs. TTX for Rab3A+/+; CON vs. TTX for Rab3A-/-, and so on. The overall means are presented ± SEM in the text, and p values are displayed above each plot, with N = the number of cells. For the characteristics of synaptic GluA2 receptor clusters and VGLUT1 sites in imaging experiments, with N = the number of dendrites (one dendrite was sampled for each cell), and statistical comparisons were made with the non-parametric Kruskal-Wallis test. Although we quantified three characteristics of the GluA2 clusters and VGLUT1 sites, size, average intensity, and integral, we have not corrected p values for multiple comparisons. Means are presented as box plots + data, with the box at 25 and 75 percentiles, whiskers at 10 and 90 percentiles, a dot indicating the mean, and a line indicating the median. In Figure 3, mEPSC amplitudes were measured within the same cell before and after NASPM application, so a paired t test was used to compare the pre- and post-NASPM mEPSC amplitudes. In Figure 5A-C the mean mEPSC amplitude for untreated and TTX-treated coverslips in the same culture are compared with a paired t test. Statistics were computed and plots were created using Origin2020 (OriginLab).